## COMP SCI 1MD3 Introduction to Programming, Winter 2018
### Douglas Stebila (Instructor), Joey Legere, Karl Knopf, Natalie Chin, Victor Chen (TAs)
### Midterm Test 1 Tuesday, Feb 13, 3:30 - 4:20 PM
### Maximum grade: 30 / 30

You must use our JupyterHub server to complete the test. Once you have completed the test, click the “Submit” button in JupyterHub.  During the test, the only program you may use is your web browser, and the only web pages you may visit are the JupyterHub site and the Python documentation on doc.python.org.  You may not use Google, StackOverflow, IDLE, your own local Jupyter, or any other website or program.  University Technology Services will be logging network access for students in this class during the test period.  Please close all other programs on your computer and all other browser tabs and windows.  During the test, please raise your hand if you need assistance, would like scrap paper, or need to go to the washroom.  The test will be over at 4:20.  If you are finished early, you may submit your notebook on JupyterHub and then quietly leave.  However, we will ask you to remain in your seat during the last 10 minutes to avoid disturbing other students. 

## Part I: Multiple Choice and Short Answer (9 points)

#### Question 1 (1 point)

Which of the following lines of code checks if the value of x is equal to the value of y?

a) `x = y`

b) `x == y`

Uncomment your answer in the cell below

```python
def question1():
    
    # Uncomment the line corresponding to your answer
    # That is, just delete the '#' in front of the first line to answer 'a', and the second line for 'b'
    
    #ans = 'a'
    
    ans = 'b'
    
    return ans
```

```python
# Run this code cell to check your answer to question 1
print("Your answer is: %s" % question1())
```

Output:
```
Your answer is: b
```

<!-- grade_id: cell-78b77369810f8c27, points: 1, grade: true, solution: false -->
```python
# Hidden tests for question 1
### BEGIN HIDDEN TESTS
assert question1() == 'b'
### END HIDDEN TESTS
```

#### Question 2 (2 points)

What types of the following types variables are mutable? In other words, what data types can you change an element of?

If the variable type is mutable, uncomment it. Otherwise, leave it commented out. Select all that apply.

```python
def question2():
    
    ## If a variable below is mutable, uncomment it. Otherwise, leave it commented out.
    
    #type1 = 'strings'
    type2 = 'lists'
    type3 = 'dictionaries'
    #type4 = 'tuples'
    
    
    ## Do not modify any of the code below this line
    mutable_variables = []
    for i in [1, 2, 3, 4]:
        
        if 'type%s'%i in locals():
            mutable_variables.append(locals()['type%s' % i])
            
    return mutable_variables
```

```python
# Use this cell to verify your answer
print("You have indicated that the following variables are mutable:", question2())
```

Output:
```
You have indicated that the following variables are mutable: ['lists', 'dictionaries']
```

<!-- grade_id: cell-05a2366b6609ef2b, points: 1, grade: true, solution: false -->
```python
# Hidden tests for question 2
### BEGIN HIDDEN TESTS
assert 'strings' not in question2()
assert 'tuples' not in question2()
### END HIDDEN TESTS
```

<!-- grade_id: cell-8d9eefa3f460b5b2, points: 1, grade: true, solution: false -->
```python
# Hidden tests for question 2
### BEGIN HIDDEN TESTS
assert 'lists' in question2()
assert 'dictionaries' in question2()
### END HIDDEN TESTS
```

#### Question 3 (2 points)

Given a list containing 7 items, (for instance `l = ['a', 'b', 'c', 'd', 'e', 'f', 'g']`), which of the following code segments will return the 3rd element (`'c'`) up to **and including** the 6th element (`'f'`).  Select all that apply.

a) `l[3:6]`

b) `l[2:6]`

c) `l[3:-1]`

d) `l[2:-1]`

e) `l[-5:-1]`

f) `l[-5:7]`

```python
def question3():
    
    ## Uncomment all letters corresponding to what you believe are the CORRECT code snippets
    
    #a = True
    b = True
    #c = True
    d = True
    e = True
    #f = True
    
    ## Do not modify the code below
    answers = []
    for i in 'abcdef':
        
        if i in locals(): answers.append(i)
            
    return answers
```

```python
# Use this code cell to verify your answer
print("You have selected: ", question3())
```

Output:
```
You have selected:  ['b', 'd', 'e']
```

<!-- grade_id: cell-6a41b0baa1f9691b, points: 1, grade: true, solution: false -->
```python
# Hidden tests for question 3
### BEGIN HIDDEN TESTS
assert 'a' not in question3()
assert 'c' not in question3()
assert 'f' not in question3()
### END HIDDEN TESTS
```

<!-- grade_id: cell-5a96c613037d52a2, points: 1, grade: true, solution: false -->
```python
# Hidden tests for question 3
### BEGIN HIDDEN TESTS
assert 'b' in question3()
assert 'd' in question3()
assert 'e' in question3()
### END HIDDEN TESTS
```

#### Question 4 (Short answer, 2 points)

Let's say you want to write a program that keeps track of all of the movies you've watched, and the number of times you have watched them. What sort of data structure would you use, and why?

<!-- grade_id: cell-6deddb6a2c4ba9f6, points: 2, grade: true, solution: true -->
```python
# Write your answer here
Dictionary. Keys are movie name, values are number of times watched. Order doesn't matter, so dictionaries are a good choice.
```

#### Question 5 (Short answer, 2 points)

There is a problem with the code below. What is it, and how would you fix it? (The code is supposed to add all the numbers from 1 to N)

~~~
def counting(N):    
    s = 0
    while i < N:
        s = s + i
    i = i + 1
    return s
~~~

<!-- grade_id: cell-11c5d168ed504377, points: 2, grade: true, solution: true -->
```python
# Write your answer here
- i is not initialized: add line i = 1
- infinite loop: i = i + 1 should be indented to be inside the while loop
- off-by-one error: only adds up to N-1, but not including N: change < to <=
```

## Part II: Programming Questions (21 points total)

#### Question 6: Wizarding Currencies (10 points)

In the wizarding world, money comes in three denominations: bronze Knuts, silver Sickles, and golden Galleons.

29 Knuts make one Sickle, and 17 Sickles make one Galleon.

Write two functions (worth 5 points each): 

* `knuts(k, s, g)` that takes in a number of Knuts, Sickles, and Galleons, and returns an integer representing the total number of Knuts the coins are worth together.

* `ksg(k)`, which, given a number of Knuts, finds their value in Knuts, Sickles, and Galleons. Return your value as a dictionary, where the key represents the type of coin (Knut, Sickle, or Galleon), and the values are the number of that type of coin. See the test cases for an example.

<!-- grade_id: cell-1f4f34250fdb3b4c, points: 0, grade: false, solution: true -->
```python
def knuts(k, s, g):
    
    # Your code here.
    
    ### BEGIN SOLUTION
    return k + s*29 + g*29*17
    ### END SOLUTION
```

```python
knuts(1, 1, 1) == 523
```

Output:
```
True
```

```python
knuts(33, 10, 1) == 816
```

Output:
```
True
```

<!-- grade_id: cell-5ad04bd63af3f9c3, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for knuts()
### BEGIN HIDDEN TESTS
assert knuts(100, 100, 100) == 52300
### END HIDDEN TESTS
```

<!-- grade_id: cell-f1965df1fed0c003, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for knuts()
### BEGIN HIDDEN TESTS
assert knuts(1, 0, 1) == 494
### END HIDDEN TESTS
```

<!-- grade_id: cell-98e6f34c9a37a4ba, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for knuts()
### BEGIN HIDDEN TESTS
assert knuts(1, 1, 0) == 30
### END HIDDEN TESTS
```

<!-- grade_id: cell-68879e5ba75507f0, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for knuts()
### BEGIN HIDDEN TESTS
assert knuts(10, 20, 30) == 15380
### END HIDDEN TESTS
```

<!-- grade_id: cell-d075af1caf6804cb, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for knuts()
### BEGIN HIDDEN TESTS
assert knuts(1, 29, 0) == 842
### END HIDDEN TESTS
```

<!-- grade_id: cell-d7c92a5c665ba558, points: 0, grade: false, solution: true -->
```python
def ksg(k):
    
    # Your code here.
    
    ### BEGIN SOLUTION
    g = k // 493
    s = (k%493) // 29
    k = k % 29
    
    return {"Knuts": k, "Sickles": s, "Galleons": g}
    ### END SOLUTION
    
```

```python
ksg(523) == {'Galleons': 1, 'Knuts': 1, 'Sickles': 1}
```

Output:
```
True
```

```python
ksg(30) == {'Galleons': 0, 'Knuts': 1, 'Sickles': 1}
```

Output:
```
True
```

```python
ksg(4096) == {'Galleons': 8, 'Knuts': 7, 'Sickles': 5}
```

Output:
```
True
```

<!-- grade_id: cell-de24a24140239f3e, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for ksg()
### BEGIN HIDDEN TESTS
assert ksg(1) == {'Galleons': 0, 'Knuts': 1, 'Sickles': 0}
### END HIDDEN TESTS
```

<!-- grade_id: cell-19dcb1a03dfbe7b5, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for ksg()
### BEGIN HIDDEN TESTS
assert ksg(29) == {'Galleons': 0, 'Knuts': 0, 'Sickles': 1}
### END HIDDEN TESTS
```

<!-- grade_id: cell-7601faea24544373, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for ksg()
### BEGIN HIDDEN TESTS
assert ksg(493) == {'Galleons': 1, 'Knuts': 0, 'Sickles': 0}
### END HIDDEN TESTS
```

<!-- grade_id: cell-f19d80b0ee96e091, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for ksg()
### BEGIN HIDDEN TESTS
assert ksg(987) == {'Galleons': 2, 'Knuts': 1, 'Sickles': 0}
### END HIDDEN TESTS
```

<!-- grade_id: cell-2132b8b13ffd6c55, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for ksg()
### BEGIN HIDDEN TESTS
assert ksg(493*5 + 29*12 + 28) == {'Galleons': 5, 'Knuts': 28, 'Sickles': 12}
### END HIDDEN TESTS
```

#### Question 7a: Checking Primes (5 points)

Write a function, `isPrime(N)`, that returns `True` if `N` is a prime number, and `False` otherwise. You may assume that `N` will always be positive, and greater than 1.

A number is prime if it is only divisible by one and itself. A simple way to check if a number is prime is to check the remainder after dividing by every positive number (other than 1) below `N`. If any remainder is equal to zero, the number is not prime.

<!-- grade_id: cell-ffec72cc583c9c6c, points: 0, grade: false, solution: true -->
```python
# Write your answer here.

### BEGIN SOLUTION

def isPrime(N):
    
    intsqrt = int(N**0.5) + 1    
    for i in range(2, intsqrt):
        
        if N % i == 0:
            return False
        
    return True

### END SOLUTION

```

```python
isPrime(4) == False
```

Output:
```
True
```

```python
isPrime(13) == True
```

Output:
```
True
```

```python
isPrime(101) == True
```

Output:
```
True
```

```python
isPrime(1993) == True
```

Output:
```
True
```

```python
isPrime(1991) == False
```

Output:
```
True
```

<!-- grade_id: cell-e3e4fb9cd0d5cda2, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for isPrime()
### BEGIN HIDDEN TESTS
assert isPrime(123457)
### END HIDDEN TESTS
```

<!-- grade_id: cell-26c09ac2650e01bc, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for isPrime()
### BEGIN HIDDEN TESTS
assert not isPrime(33)
### END HIDDEN TESTS
```

<!-- grade_id: cell-945ed1173ef1217d, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for isPrime()
### BEGIN HIDDEN TESTS
assert isPrime(59)
### END HIDDEN TESTS
```

<!-- grade_id: cell-31c9848c88d2fc07, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for isPrime()
### BEGIN HIDDEN TESTS
assert not isPrime(24)
### END HIDDEN TESTS
```

<!-- grade_id: cell-80c43bf83b599c6e, points: 1, grade: true, solution: false -->
```python
# Hidden test cases for isPrime()
### BEGIN HIDDEN TESTS
assert isPrime(337)
### END HIDDEN TESTS
```

#### Question 7b: Find the Next Prime (3 points)

Write a function, `nextPrime(N)` that finds the next number larger than `N` that is prime. You can rest assured that there is always another prime number above `N`.

<!-- grade_id: cell-bc772084dd8eaf72, points: 0, grade: false, solution: true -->
```python
# Write your answer here.

### BEGIN SOLUTION

def isPrime(N):
    
    intsqrt = int(N**0.5) + 1    
    for i in range(2, intsqrt):
        
        if N % i == 0:
            return False
        
    return True

def nextPrime(N):
    
    i = N+1
    
    while not isPrime(i): i += 1
        
    return i

### END SOLUTION
```

```python
nextPrime(12) == 13
```

Output:
```
True
```

```python
nextPrime(13) == 17
```

Output:
```
True
```

```python
nextPrime(2**16) == 65537
```

Output:
```
True
```

<!-- grade_id: cell-3b30499e72fbe01b, points: 1, grade: true, solution: false -->
```python
# Hidden tests for nextPrime()
### BEGIN HIDDEN TESTS
assert nextPrime(2**10) == 1031
### END HIDDEN TESTS
```

<!-- grade_id: cell-492d87ff3ec94583, points: 1, grade: true, solution: false -->
```python
# Hidden tests for nextPrime()
### BEGIN HIDDEN TESTS
assert nextPrime(2**13) == 8209
### END HIDDEN TESTS
```

<!-- grade_id: cell-0bb858edb4f165e7, points: 1, grade: true, solution: false -->
```python
# Hidden tests for nextPrime()
### BEGIN HIDDEN TESTS
assert nextPrime(10**6) == 1000003
### END HIDDEN TESTS
```

#### Question 8a: Largest Product (2 points)

In the following number, the four adjacent digits that give the largest product are:
9 × 9 × 8 × 9 = 5832.

5100474821663704844031**9989**00088952434506585412275886

Write a function `max_product` that takes as input a base 10 number `s` (represented as a string), and outputs the product of the **twelve** adjacent digits that have the greatest product.

<!-- grade_id: cell-9cc76527702ece7b, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
    
### BEGIN SOLUTION
def max_product(s):
    MAX_PRODUCT = 0
    for i in range(len(s)-11):

        prod = 1

        for j in range(12): prod *= int(s[i+j])

        if prod > MAX_PRODUCT:
            MAX_PRODUCT = prod
        

    
    return MAX_PRODUCT

### END SOLUTION
```

```python
max_product('123456789123456789') == 182891520
```

Output:
```
True
```

<!-- grade_id: cell-ab5bfc6fa8d874ac, points: 2, grade: true, solution: false -->
```python
# Hidden test cell for max_product()
### BEGIN HIDDEN TESTS
s = '7316717653133062491922511967442657474235534919493496983520312774506326239578318016984801869478851843858615607891129494954595017379583319528532088055111254069874715852386305071569329096329522744304355766896648950445244523161731856403098711121722383113622298934233803081353362766142828064444866452387493035890729629049156044077239071381051585930796086670172427121883998797908792274921901699720888093776657273330010533678812202354218097512545405947522435258490771167055601360483958644670632441572215539753697817977846174064955149290862569321978468622482839722413756570560574902614079729686524145351004748216637048440319989000889524345065854122758866688116427171479924442928230863465674813919123162824586178664583591245665294765456828489128831426076900422421902267105562632111110937054421750694165896040807198403850962455444362981230987879927244284909188845801561660979191338754992005240636899125607176060588611646710940507754100225698315520005593572972571636269561882670428252483600823257530420752963450'
assert max_product(s) == 4702924800
### END HIDDEN TESTS
```

#### Question 8b (1 point)

Write a function `max_product2` that finds the value of the second largest product of twelve adjacent digits in a given number.

<!-- grade_id: cell-b9d4509836dadcaa, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
    
### BEGIN SOLUTION
def max_product2(s):
    MAX_PRODUCT = 0
    MAX_PRODUCT2 = 0
    for i in range(len(s)-11):

        prod = 1

        for j in range(12): prod *= int(s[i+j])

        if prod > MAX_PRODUCT:
            MAX_PRODUCT2 = MAX_PRODUCT
            MAX_PRODUCT = prod
        
    return MAX_PRODUCT2
### END SOLUTION
```

```python
max_product2('123456789123456789') == 121927680
```

Output:
```
True
```

<!-- grade_id: cell-3e1c72929b24b413, points: 1, grade: true, solution: false -->
```python
# Hidden test for max_product2()
### BEGIN HIDDEN TESTS
s = '7316717653133062491922511967442657474235534919493496983520312774506326239578318016984801869478851843858615607891129494954595017379583319528532088055111254069874715852386305071569329096329522744304355766896648950445244523161731856403098711121722383113622298934233803081353362766142828064444866452387493035890729629049156044077239071381051585930796086670172427121883998797908792274921901699720888093776657273330010533678812202354218097512545405947522435258490771167055601360483958644670632441572215539753697817977846174064955149290862569321978468622482839722413756570560574902614079729686524145351004748216637048440319989000889524345065854122758866688116427171479924442928230863465674813919123162824586178664583591245665294765456828489128831426076900422421902267105562632111110937054421750694165896040807198403850962455444362981230987879927244284909188845801561660979191338754992005240636899125607176060588611646710940507754100225698315520005593572972571636269561882670428252483600823257530420752963450'
max_product2(s) == 1567641600
### END HIDDEN TESTS
```

Output:
```
True
```