## COMP SCI 1MD3 Introduction to Programming, Winter 2018
### Douglas Stebila (Instructor), Joey Legere, Karl Knopf, Natalie Chin, Victor Chen (TAs)
### Lab 6 Assigned Wednesday, Feb 21, Due Wednesday Feb 28, 5pm
### Maximum grade: 20 / 16

The purpose of this lab is to:
* Learn to check data types before manipulating data
* Practice handling different types of exceptions gracefully without your program breaking

#### Practice Question 1: Finding a Product

Consider the following function that multiplies all the numbers in a given list:

```python
def prod(l):
    
    p = 1
    
    for i in l:
        
        p *= i
        
    return p
```

This works quite well, given a list of all integers, or a mix with floating point numbers:

```python
prod([1,2,3,7]) == 42
```

```python
prod([1.0, 2.0, 3.1, 17.5]) == 108.5
```

However, adding strings or lists produces unexpected behaviours:

```python
prod([1, 2, 3, 'badger '])
```

```python
prod([2, [7, 8, 9], 3])
```

```python
prod([1, 3, 7, '', []])
```

We can use the `type` function to get the type of a variable and make use of that in our code.  For example:

```python
x = 17
if type(x) in (int, float):
    print("x is a number")
elif type(x) == str:
    print("x is a string")
elif type(x) in (list, set, tuple):
    print("x is a list/set/tuple")
elif type(x) == dict:
    print("x is a dictionary")
else:
    print("x is a " + str(type(x)))
```

Write a function, `prodBetter(l)` that performs the same function as `prod(l)`, but ignores any elements of `l` that are not integers or floating point numbers.

```python
# Your code here
```

```python
prodBetter([2, [7, 8, 9], 3, 'kittens']) == 6
```

#### Practice Question 2: prodEvenBetter()

Let's build on the `prodBetter(l)` function. Say that '3.14' is an element of the list. We want our function to be able to multiply by this number. However, we obviously don't want to convert the string 'kittens' to a number (<a href='https://xkcd.com/26/'>generally one should not perform mathematics on cats </a>).

Improve the function so that it tries to convert any strings in the list into floating point numbers, but skips any string that cannot be converted to a number.  You will have to make use of try/except to catch any conversion errors that occur.

```python
# Write prodEvenBetter here
```

```python
prodEvenBetter([2, [7, 8, 9], 3, 'kittens']) == 6
```

```python
prodEvenBetter([2, [7, 8, 9], 3, '3.14']) == 18.84
```

#### Practice Question 3: prodBest(l)

What a terrific function. Too bad it can't work with hexadecimal numbers (yet).

Create a function `prodBest(l)` that does everything `prodBetter(l)` does, but also tries to convert any strings that are provided into hexadecimal.

Modify the given function hex2dec() so that it raises a `BaseException` if it cannot convert the given number from hex. `prodBest()` should skip any strings in the list that raise this particular error.



```python
# Modify to raise a BaseException if it cannot convert a given string
def hex2dec(x):

    return int(x, 16)
        
```

```python
# Write prodBest here.
```

```python
prodBest(['1A', 'F']) == 390
```

```python
prodBest([[], 3.0, '3.14', 'cats', '1a']) == 244.92
```

#### Question 1 (4 points):  Adding Numbers (But Better!)

Write a function, `addBetter(x, y)` that takes two numbers, adds them, and **returns a floating point number (even if all the inputs are integers)**. Your function should try to convert strings to floats, and handle any weird inputs gracefully. If you are unable to add `x` and `y`, return `False`.

```python
# Use this code cell to run your tests
```

<!-- grade_id: cell-ffec72cc583c9c6c, points: 0, grade: false, solution: true -->
```python
# Write your answer here.

###BEGIN SOLUTION

def addBetter(x, y):
    
    if type(x) not in (int, float, str) or type(y) not in (int, float, str):
        return False
    
    else:
        
        try:
            x = float(x)
            y = float(y)
        except:
            return False
        
        return x + y

### END SOLUTION
```

```python
addBetter(3, 4) == 7
```

```python
addBetter('3.14', '2.72') == 5.86
```

```python
addBetter('cat', 3) == False
```

```python
type(addBetter(0, 0)) == float
```

<!-- grade_id: cell-31c9848c88d2fc07, points: 1, grade: true, solution: false -->
```python
# Hidden test for addBetter
### BEGIN HIDDEN TESTS
assert type(addBetter('3', 1)) == float
### END HIDDEN TESTS
```

<!-- grade_id: cell-ca28d218ed96521f, points: 1, grade: true, solution: false -->
```python
# Hidden test for addBetter
### BEGIN HIDDEN TESTS
assert not addBetter('potato', 'salad') 
### END HIDDEN TESTS
```

<!-- grade_id: cell-dc1492c7ee471b53, points: 1, grade: true, solution: false -->
```python
# Hidden test for addBetter
### BEGIN HIDDEN TESTS
assert addBetter('33.89', 5.13) == 39.02
### END HIDDEN TESTS
```

<!-- grade_id: cell-4b1c513d2ecefb29, points: 1, grade: true, solution: false -->
```python
# Hidden test for addBetter
### BEGIN HIDDEN TESTS
assert (addBetter('3', '-7.1')) == -4.1
### END HIDDEN TESTS
```

#### Question 2 (4 points): Comparing Rainy Seasons

Write a function, `precipiation(d, a, b)` that takes parameters:

`d`: Dictionary mapping a season to an amount of precipitation for that season.<br>
`a`: One season<br>
`b`: Another season

Your function should return `True` if the rainfall in season `a` was greater than it was in season `b`, and `False` otherwise. If the rainfall was equal, return `False`.

Your function should handle the following errors the following ways:

If either a or b is not a valid season in the dictionary, you should return an error message (as a string) formatted **exactly** as:<br>
"Error: Key is not in Dictionary."

If the values for the rainfall in each season are not valid for comparison, that is, one of them is not a number, return an error message (as a string) formatted **exactly** as:<br>
"Error: Invalid data types in Dictionary?"

You may assume that no test case will produce both types of errors (so you do not need to decide whether or not to return a Key Error or Invalid Type error when you encounter both). However, you should not make any assumptions about the seasons that exist--sometimes there may be a 'potato' season.

```python
# Use this code cell to run your tests
```

<!-- grade_id: cell-bc772084dd8eaf72, points: 0, grade: false, solution: true -->
```python
# Write your answer here.

### BEGIN SOLUTION

def precipitation(d, a, b):
    
    try:
        if type(d[a]) not in (int, float) or type(d[b]) not in (int, float):
            return "Error: Invalid data types in Dictionary?"
    except KeyError: return "Error: Key is not in Dictionary."
        
    return d[a] > d[b]

### END SOLUTION


```

```python
precipitation({"Spring": 813, "Summer":337, "Fall": 581, "Winter": 468}, "Spring", "Summer") == True
```

```python
precipitation({"Spring": 982, "Summer":484, "Fall": 557, "Winter": 313}, "Winter", "Spring") == False
```

```python
precipitation({"Spring": 887, "Summer":529, "Fall": 616, "Winter": 286}, "Winter", "Potato") == "Error: Key is not in Dictionary."
```

```python
precipitation({"Spring": 887, "Summer":529, "Fall": 616, "Winter": 286, "Potato": 17}, "Winter", "Potato") == True
```

```python
precipitation({"Spring": 'cats', "Summer": 313, "Fall": 1457, "Winter": 354}, "Summer", "Spring") == 'Error: Invalid data types in Dictionary?'
```

<!-- grade_id: cell-3b30499e72fbe01b, points: 1, grade: true, solution: false -->
```python
# Hidden test cell for precipitation
### BEGIN HIDDEN TESTS
assert precipitation({"Spring": 'cats', "Summer": 'dogs', "Fall": 1457, "Winter": 354}, "Summer", "Spring") == 'Error: Invalid data types in Dictionary?'
### END HIDDEN TESTS
```

<!-- grade_id: cell-beb43870c366f25b, points: 1, grade: true, solution: false -->
```python
# Hidden test cell for precipitation
### BEGIN HIDDEN TESTS
assert not precipitation({"potato1":33, "potato2": 47}, "potato1", "potato2")
### END HIDDEN TESTS
```

<!-- grade_id: cell-e1eebcbf3d5864b6, points: 1, grade: true, solution: false -->
```python
# Hidden test cell for precipitation
### BEGIN HIDDEN TESTS
assert precipitation({"Spring": 'cats', "Summer": 313, "Fall": 1457, "Winter": 354}, "cat", "dog") == 'Error: Key is not in Dictionary.'
### END HIDDEN TESTS
```

<!-- grade_id: cell-87385834f481adc1, points: 1, grade: true, solution: false -->
```python
# Hidden test cell for precipitation
### BEGIN HIDDEN TESTS
assert precipitation({"Spring": 'cats', "Summer": 'dogs', "Fall": 1457, "Winter": 354}, "Fall", "Winter")
### END HIDDEN TESTS
```

#### Question 3 (4 points): Checking Valid Inputs

Write a function `valid(f, l)` that, given a function `f` and a list of potential inputs `l`, returns a list, tuple, or set of the elements of `l` that are valid inputs to `f`.  How will you know if something is a valid input to `f`?  If you call `f` on the element and an exception is thrown, you know it was not a valid input--review practice question 3 where we try to use the hex2dec() function, but ignore any error we get if the input was not valid.

For instance, if `f` is the `abs()` function, it can only accept integers and floating point numbers. If it is the `max()` function, it can only accept lists or tuples, if it is the `sqrt()` function, it can only accept non-negative numbers, etc...

<!-- grade_id: cell-6e59a5187b224404, points: 0, grade: false, solution: true -->
```python
# Write your answer here.

### BEGIN SOLUTION

def valid(f, l):
    
    good = []
    
    for i in l:
        
        try:
            f(i)
            good.append(i)
        except:
            pass
        
    return good

### END SOLUTION

```

```python
from math import sqrt
valid(sqrt, [-3, 0, 5]) == [0, 5]
```

```python
from math import asin
valid(asin, [0.3, -0.1, -1, 1, 'cat', '1', 1.3]) == [0.3, -0.1, -1, 1]
```

```python
valid(chr, [-1, 10, 123.3, 2**21]) == [10]
```

```python
def onlyStrings(x):
    
    if type(x) == str:
        return True
    else:
        raise UserWarning

valid(onlyStrings, ['cat', 123, [1, 2, 3]]) == ['cat']
```

<!-- grade_id: cell-d6df13b487b53f1f, points: 1, grade: true, solution: false -->
```python
# Hidden test cell for valid()
### BEGIN HIDDEN TESTS
assert valid(pow, [(1,2), 'honeybadger', 3, 5.14, {'potato':'tomato'}]) == []
### END HIDDEN TESTS
```

<!-- grade_id: cell-0dc4921f996ca83d, points: 1, grade: true, solution: false -->
```python
# Hidden test cell for valid()
### BEGIN HIDDEN TESTS
assert valid(int, ['1', '3.14', '-5', {}]) == ['1', '-5']
### END HIDDEN TESTS
```

<!-- grade_id: cell-d9a6d490ab954561, points: 1, grade: true, solution: false -->
```python
# Hidden test cell for valid()
### BEGIN HIDDEN TESTS
assert valid(set, [(1,2), 'honeybadger', 3, 5.14, [{'potato':'tomato'}]]) ==[(1, 2), 'honeybadger']
### END HIDDEN TESTS
```

<!-- grade_id: cell-104063570791e181, points: 1, grade: true, solution: false -->
```python
# Hidden test cell for valid()
### BEGIN HIDDEN TESTS
from cmath import sqrt
assert valid(sqrt, [-3, '1', 'cat', -1.3, 25]) == [-3, -1.3, 25]
### END HIDDEN TESTS
```

#### Question 4a (2 points): Finding Edge Cases I

This program is a partially correct solution to the winning streak problem from a previous lab. Given a string, list, or tuple as input, find the longest winning streak. This solution contains an error. Find a test case that returns an incorrect value, or raises an error.

**Do not modify or try to correct the code for this question, simply find a problematic input for the function**

<!-- grade_id: cell-9209a400fa396373, points: 0, grade: false, solution: true -->
```python
# Modify the value of this variable so that winningStreakBad produces incorrect output
PROBLEM_INPUT1 = ''

# Do not modify the code below this line
def winningStreakBad(s):
    
    
    if type(s) not in (tuple, list, str):
        return 0    
    
    if len(s) == 0:
        return 0
    
    wins = 1
    maxwins = 1
    
    for i in range(len(s)):
        
        if s[i] == s[i-1]:
            wins += 1
        else:
            wins = 1
        
        if wins > maxwins:
            maxwins = wins
            
    return maxwins
        

### BEGIN SOLUTION

def winningStreakBad(s):
 
    if type(s) not in (tuple, list, str):
        return 0

    if len(s) == 0:
        return 0

    wins = 1
    maxwins = 1
    
    for i in range(len(s)):
        
        if s[i] == s[i-1]:
            wins += 1
        else:
            wins = 1
        
        if wins > maxwins:
            maxwins = wins
            
    return maxwins

def winningStreakGood(s):
    
    if type(s) not in (tuple, list, str):
        return 0
    
    if len(s) == 0:
        return 0    
    
    wins = 1
    maxwins = 1
    
    for i in range(1, len(s)):
        
        if s[i] == s[i-1]:
            wins += 1
        else:
            wins = 1
        
        if wins > maxwins:
            maxwins = wins
            
    return maxwins

### END SOLUTION
```

```python
## This is what the autograder will see:
winningStreakBad(PROBLEM_INPUT1)
```

<!-- grade_id: cell-33490ad25625a24f, points: 2, grade: true, solution: false -->
```python
# Hidden test cell for 4a
### BEGIN HIDDEN TESTS
def winningStreakBad(s):
 
    if type(s) not in (tuple, list, str):
        return 0

    if len(s) == 0:
        return 0

    wins = 1
    maxwins = 1
    
    for i in range(len(s)):
        
        if s[i] == s[i-1]:
            wins += 1
        else:
            wins = 1
        
        if wins > maxwins:
            maxwins = wins
            
    return maxwins

def winningStreakGood(s):
    
    if type(s) not in (tuple, list, str):
        return 0
    
    if len(s) == 0:
        return 0    
    
    wins = 1
    maxwins = 1
    
    for i in range(1, len(s)):
        
        if s[i] == s[i-1]:
            wins += 1
        else:
            wins = 1
        
        if wins > maxwins:
            maxwins = wins
            
    return maxwins
try:
    winningStreakBad(PROBLEM_INPUT1)
    winningStreakGood(PROBLEM_INPUT1)
    assert winningStreakBad(PROBLEM_INPUT1) != winningStreakGood(PROBLEM_INPUT1)
    
except AssertionError: # The outputs of the functions are the same
    assert False

except NameError: # They did not define PROBLEM_INPUT1
    assert False
    
except: # They found a different error
    assert True

### END HIDDEN TESTS
```

#### Question 4b (2 points): Finding Edge Cases II

The code below tries to solve the linear system:
$$ax + b = cx + d$$

Find an input (A, B, C, D) that produces incorrect results. For the purposes of this question, `False` is considered a correst result for non-numeric inputs (strings, lists, etc.)

<!-- grade_id: cell-8dcff65a305b9ef9, points: 0, grade: false, solution: true -->
```python
# Modify these values so that solveLinearBad does not produce correct results
A = 1
B = 2
C = 3
D = 4


# Do not modify the code below this line
def solveLinearBad(a, b, c, d):
    
    for i in [a, b, c, d]:
        if type(i) not in (int, float):
            return False
    
    x = (b - d) / (c - a)
    
    return x

### BEGIN SOLUTION

def solveLinearBad(a, b, c, d):
    
    for i in [a, b, c, d]:
        if type(i) not in (int, float):
            return False    
    
    x = (b - d) / (c - a)
    
    return x

### END SOLUTION

```

```python
## This is what the autograder will see:
solveLinearBad(A, B, C, D)
```

<!-- grade_id: cell-b97f26dd1961c7c2, points: 2, grade: true, solution: false -->
```python
# Hidden test cell for 4b
### BEGIN HIDDEN TESTS
def solveLinearBad(a, b, c, d):
    
    for i in [a, b, c, d]:
        if type(i) not in (int, float):
            return False    
    
    x = (b - d) / (c - a)
    
    return x

try:
    solveLinearBad(A, B, C, D)
    assert False # only gets called if solveLinearBad() ran correctly

except AssertionError: # Their function ran correctly.
    assert False

except NameError: # They did not define A, B, C, or D
    assert False
    
except: 
    assert True    
    
### END HIDDEN TESTS
```

#### Bonus (4 points): Why was 6 afraid of 7?

#### Because 7 8 9!

Actually, this is mathematically impossible. "Cannibal numbers" as they're called can only eat numbers that are smaller than them, or the same size. After eating another number, the number that gets eaten disappears, and the number that ate them increases by 1. 

For example, given the following set of numbers, let's pretend that **9** is the cannibal.<br>
[27, **9**, 3, 8, 11]

9 could start by eating 3 or 8. Let's say it eats 3, then the numbers would be:<br>
[27, **10**, 8, 11]

Then it could eat 8, so the numbers would be:<br>
[27, **11**, 8, 11]

And lastly, it could eat 11:<br>
[27, **12**, 8]

Your task is to write a program that, given a list of numbers and a target number, determines how many numbers in the list can become **exactly** that number. For instance, in our example, two numbers could reach the value of 10. As we saw, 9 took on that value in the example. However, if 8 were the cannibal, it could:

* Eat 3 to become 9
* Eat 9 to become 10

So it could also hold the value of 10.

For a target value of 11, the answer would also be 2, because 9 can become 11, and 11 is already 11.

Your function may be passed a string of values separated by commas instead of a list--1 out of 4 points will be awarded for handling this case correctly.

<!-- grade_id: cell-77063b4697eac250, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
### BEGIN SOLUTION

# One solution
def cannibal(l, target):
    
    if type(l) == str:
        l = l.split(",")
    l = [int(i) for i in l]    
    
    cannibals = 0
    for i in range(len(l)):
        
        eaten = []
        this = l[i]
        
        j = 0
        while j < len(l):
            
            if this == target:
                cannibals += 1
                break
                
            if this >= l[j] and i != j and j not in eaten:
                this += 1
                eaten.append(j)
                j = 0
            j += 1
    
    return cannibals

# Alternative solution, sorted, easier to understand?
def cannibal(l, target):
    
    if type(l) == str:
        l = l.split(",")
    l = [int(i) for i in l]
    
    cannibals = 0
    l.sort()
    
    for i in range(len(l)):
        
        this = l[i]
        
        if this == target:
            cannibals += 1
            continue
        
        for j in range(len(l)):
            
            if l[j] <= this and j != i:
                this += 1
            
            elif l[j] > this:
                break
                
            if this == target:
                cannibals += 1
                break
                
    return cannibals


### END SOLUTION
```

```python
cannibal([27, 9, 3, 8, 11], 11) == 2
```

```python
cannibal([27, 9, 3, 8, 11], 50) == 0
```

```python
cannibal([3, 3, 3, 2, 2, 2, 1, 1, 1], 4) == 9
```

```python
cannibal("1,2,3,4,5", 5) == 4
```

<!-- grade_id: cell-e58bfa8edb7ad7d4, points: 1, grade: true, solution: false -->
```python
# Hidden tests for bonus
### BEGIN HIDDEN TESTS
assert cannibal([106, 110, 107, 109, 33, 44, 10], 109) == 3
### END HIDDEN TESTS
```

<!-- grade_id: cell-75d365964471fae0, points: 1, grade: true, solution: false -->
```python
# Hidden tests for bonus
### BEGIN HIDDEN TESTS
assert cannibal([1, 1, 1, 2, 3], 5) == 5
### END HIDDEN TESTS
```

<!-- grade_id: cell-c0f9ca7e33d6c880, points: 1, grade: true, solution: false -->
```python
# Hidden tests for bonus
### BEGIN HIDDEN TESTS
assert cannibal("1,2,3,4,6", 7) == 3
### END HIDDEN TESTS
```

<!-- grade_id: cell-a55554dc00825531, points: 1, grade: true, solution: false -->
```python
# Hidden tests for bonus
### BEGIN HIDDEN TESTS
assert cannibal([1, 6, 4, 9, 2, 6, 3, 16], 15) == 1
### END HIDDEN TESTS
```