## COMP SCI 1MD3 Introduction to Programming, Winter 2018
### Douglas Stebila (Instructor), Joey Legere, Karl Knopf, Natalie Chin, Victor Chen (TAs)
### Lab 7 Assigned Wednesday, Feb 28, Due Wednesday March 7, 5pm
### Maximum grade: 20 / 16

The purpose of this lab is to:
* Practice extending existing classes
* Implement basic operators for classes such as +, -, `__getitem__` and *
* Create your own custom classes and exceptions

#### Practice Question: Blackjack

Write a program that draws cards from a deck, and simulates a game of blackjack.

First, create a class representing a card by filling in the template below. For the value method, Jacks, Kings, and Queens are worth 10, and aces should be returned as 1--whether or not an ace is high (11) or low (1) can be decided by the program that is "playing" blackjack later.

```python
class Card:
    
    def __init__(self, suit, face):
        
        pass # Your code here.
        
    def value(self):
        '''
        Returns the value of the card as an integer.
        For Aces, return 1.
        '''
        
        pass # Your code here.
```

```python
c1 = Card("Spades", 3)
c1.value() == 3
```

```python
c2 = Card("Hearts", "K")
c2.value() == 10
```

```python
c3 = Card("Diamonds", "A")
c3.value() == 1
```

Now, create a `Deck` class that consists of 52 cards by filling in the template below. Represent the suit of cards as a string: `("Spades", "Diamonds", "Hearts", "Clubs")` and the face of cards as a single character or integer: `[2, 3, 4, 5, 6, 7, 8, 9, 10, "J", "Q", "K", "A"]`

```python
class Deck:
    
    def __init__(self):

            
    def draw(self):
        """
        Returns the card at the top of the deck, and removes it from the deck
        """
    
        pass
    def shuffle(self):
        """
        Shuffle the remaining cards in the deck.
        """
        pass
    
    def reshuffle(self):
        """
        Put all the cards back in the deck, and shuffle it.
        """
        pass

```

```python
d = Deck()
```

```python
card = d.draw()
print(card.suit)
print(card.face)
print(card.value()) # Pay attention to the brackets here. value() is a method, so we need to call it like one
```

```python
d.shuffle()
```

```python
card = d.draw()
print(card.suit)
print(card.face)
print(card.value()) # Pay attention to the brackets here. value() is a method, so we need to call it like one
```

The following function returns the value of a hand (as a list of `Card` objects) in blackjack:

```python
def value(hand):
    
    val = 0
    num_aces = 0
    
    for card in hand:
        this_val = card.value()
        
        if this_val != 1:
            val += this_val
        else:
            num_aces += 1
    
    while num_aces > 0:
        
        if val + 11 > 21:
            val += 1
        else:
            val += 11
            
        num_aces -= 1
    
    return val
```

#### Question 1a: Extending Geometric Shapes (2 points)

Add the method `area` to the classes `Line` and `Rectangle` that returns the area of the given shapes. The area of a line is always `0`.

<!-- grade_id: cell-a5f50ca21561ee40, points: 0, grade: false, solution: true -->
```python
class Shape:
    def setOrigin(self, x, y):
        self.x, self.y = x, y
        
    def move(self, dx, dy):
        """Move by dx and dy"""
        self.x, self.y = self.x + dx, self.y + dy

class Line(Shape):
    def __init__(self, x, y, u, v):
        """Line with start point (x, y) and end point (u, v)"""
        Shape.setOrigin(self, x, y)
        self.u, self.v = u, v
        
    def move(self, dx, dy):
        Shape.move(self, dx, dy)
        self.u, self.v = self.u + dx, self.v + dy
        
    def area(self):
        pass # Replace this line with your body
    
    def __str__(self):
        return "Line from (" + str(self.x) + ", " + str(self.y) + \
               ") to (" + str(self.u) + ", " + str(self.v) + ")"

class Rectangle(Shape):
    def __init__(self, x, y, w, h):
        """Rectangle at (x, y) with width w and height h"""
        Shape.setOrigin(self, x, y)
        self.w, self.h = w, h
        
    def area(self):
        pass # Replace this line with your body
        
    def __str__(self):
        return "Rectangle at (" + str(self.x) + ", " + str(self.y) + \
               "), width " + str(self.w) + ", height " + str(self.h)
    
### BEGIN SOLUTION

class Shape:
    def setOrigin(self, x, y):
        self.x, self.y = x, y
    def move(self, dx, dy):
        """Move by dx and dy"""
        self.x, self.y = self.x + dx, self.y + dy

class Line(Shape):
    def __init__(self, x, y, u, v):
        """Line with start point (x, y) and end point (u, v)"""
        Shape.setOrigin(self, x, y)
        self.u, self.v = u, v
        
    def move(self, dx, dy):
        Shape.move(self, dx, dy)
        self.u, self.v = self.u + dx, self.v + dy
        
    def area(self):
        return 0
    
    def __str__(self):
        return "Line from (" + str(self.x) + ", " + str(self.y) + \
               ") to (" + str(self.u) + ", " + str(self.v) + ")"

class Rectangle(Shape):
    def __init__(self, x, y, w, h):
        """Rectangle at (x, y) with width w and height h"""
        Shape.setOrigin(self, x, y)
        self.w, self.h = w, h
        
    def area(self):
        return self.w * self.h
    
    def __str__(self):
        return "Rectangle at (" + str(self.x) + ", " + str(self.y) + \
               "), width " + str(self.w) + ", height " + str(self.h)

### END SOLUTION    
```

```python
r = Rectangle(-2, -3, 2, 3)
```

```python
l = Line(1, 2, 5, 6)
```

```python
str(r) == "Rectangle at (-2, -3), width 2, height 3"
```

```python
str(l) == "Line from (1, 2) to (5, 6)"
```

```python
l.area() == 0
```

```python
r.area() == 6
```

<!-- grade_id: cell-6f60efc3ada4cd3a, points: 1, grade: true, solution: false -->
```python
# Hidden test for 1a
### BEGIN HIDDEN TESTS
r = Rectangle(-8, 0, 2.5, 3.7)
assert r.area() == 9.25

### END HIDDEN TESTS
```

<!-- grade_id: cell-e5f9d3ac468b7100, points: 1, grade: true, solution: false -->
```python
# Hidden test for 1a
### BEGIN HIDDEN TESTS
l = Line(1, 2, 31, 96)
assert l.area() == 0
### END HIDDEN TESTS
```

#### Question 1b: Groups of Shapes (2 points)

The class below represents a group of shapes. For instance, a group could be a square, a rectangle, and a triangle. Implement the `area` method for the `Group` class.

The area of a group is just the total area of the shapes in the group. For the purpose of this assignment, we will not consider cases where shapes overlap, you simply need to add the areas of the shapes in a group independently.

<!-- grade_id: cell-d8bd6680a06909ab, points: 0, grade: false, solution: true -->
```python
class Group:
    def __init__(self):
        self.c = set()
    def add(self, s):
        """Add shape s to group"""
        self.c.add(s)
    def move(self, dx, dy):
        """Move by dx and dy"""
        for s in self.c:
            s.move(dx, dy)
            
    def area(self):
        pass # Replace this line with your body
        
    def __str__(self):
        r = "Group with:"
        for s in self.c:
            r = r + "\n  " + str(s)
        return r
    
    
### BEGIN SOLUTION
class Group:
    def __init__(self):
        self.c = set()
    def add(self, s):
        """Add shape s to group"""
        self.c.add(s)
    def move(self, dx, dy):
        """Move by dx and dy"""
        for s in self.c:
            s.move(dx, dy)
    def area(self):
        # Replace this line with your body
        return sum(s.area() for s in self.c)
    def __str__(self):
        r = "Group with:"
        for s in self.c:
            r = r + "\n  " + str(s)
        return r#
    
### END SOLUTION
```

The following cell should create a new group, and add both our line and rectangle from above to it. The output should look like:
```
Group with:
  Line from (1, 2) to (5, 6)
  Rectangle at (-2, -3), width 2, height 3
```

```python
g1 = Group()
g1.add(l)
g1.add(r)
print(g1)
```

Now, if we move every element in the group up one unit and to the right one unit, the output should look like:
```
Group with:
  Line from (2, 3) to (6, 7)
  Rectangle at (-1, -2), width 2, height 3
```

```python
g1.move(1, 1)
print(g1)
```

```python
g1.area() == 6
```

<!-- grade_id: cell-8d57d815985b12c1, points: 1, grade: true, solution: false -->
```python
# Hidden test for 1b
### BEGIN HIDDEN TESTS
r1 = Rectangle(3,3,31,3)
r2 = Rectangle(-1,7,22,44)
g = Group()
g.add(r1)
g.add(r2)
assert g.area() == 1061
### END HIDDEN TESTS
```

<!-- grade_id: cell-6b7a02039a4e74df, points: 1, grade: true, solution: false -->
```python
# Hidden test for 1b
### BEGIN HIDDEN TESTS
r1 = Rectangle(3,3,1,3.3)
r2 = Rectangle(-1,7,4,3.8)
l1 = Line(1,2,3,4)
g = Group()
g.add(r1)
g.add(r2)
g.add(l1)
assert g.area() == 18.5
### END HIDDEN TESTS
```

#### Question 1c: Creating New Shapes (4 points)

Create two new shapes: `Circle` and `Polygon`.

Classes `Circle` and `Polygon` should inherit from `Shape`. For circles, at creation the radius has to be specified and the `__str__` method should return a string in the same style as for `Line` and `Rectangle`. The are of circles should be computed in the usual way. For this, import `pi` from `math`.  The area of a non-self-intersecting polygon with $n$ vertices is defined as

> $A = {1 \over 2} \sum_{i=0}^{n-1}(x_i y_{i+1}-x_{i+1} y_i)$

where each $x_i$, $y_i$ represents a vertex (“corner”) of a polygon and where $x_n = x_0$ and $y_n = y_0$. Note that the area of a convex polygon is positive if the points are in counterclockwise order and negative if they are in clockwise order. (The Wikipedia page takes the absolute value to avoid negative areas. Don’t do this here.)

<!-- grade_id: cell-9aba98829f8fe16b, points: 0, grade: false, solution: true -->
```python
# Create classes Circle and Polygon here

class Circle(Shape):
    pass # Replace this line with your methods
    
class Polygon(Shape):
    pass # Replace this line with your methods    
    
    
### BEGIN SOLUTION
class Circle(Shape):
    def __init__(self, x, y, r):
        """Circle (x, y) with radius r"""
        Shape.setOrigin(self, x, y)
        self.r = r
    def area(self):
        from math import pi
        return self.r * self.r * pi
    def __str__(self):
        return "Circle at (" + str(self.x) + ", " + str(self.y) + \
               "), radius " + str(self.r)

class Polygon(Shape):
    def __init__(self, l):
        """Polgon with non-empty list l of vertices"""
        Shape.setOrigin(self, l[0][0], l[0][1])
        self.l = l
    def move(self, dx, dy):
        Shape.move(self, dx, dy)
        self.l = [(x + dx, y + dy) for (x, y) in self.l]
    def area(self):
        r, a = 0.0, self.l[- 1]
        for b in self.l:
            r, a = r + a[0] * b[1] - b[0] * a[1], b
        return r / 2
    def __str__(self):
        return "Polygon with " + str(self.l)
### END SOLUTION
```

```python
g2 = Group()

p = Polygon([(0, 0), (4, 0), (2, 2), (4, 4), (0, 4), (2, 2)])
g2.add(p)

c = Circle(1, -5, 1)
g2.add(c)
```

```python
str(c) == "Circle at (1, -5), radius 1"
```

```python
str(p) == "Polygon with [(0, 0), (4, 0), (2, 2), (4, 4), (0, 4), (2, 2)]"
```

```python
c.area() == 3.141592653589793
```

```python
p.area() == 8.0
```

```python
round(g2.area(), 3) == 11.142
```

<!-- grade_id: cell-8eb29d4032a54dc9, points: 1, grade: true, solution: false -->
```python
# Hidden test for 1c
### BEGIN HIDDEN TESTS
c = Circle(2, 5, 3.3)
assert round(c.area(), 4) == 34.2119
### END HIDDEN TESTS
```

<!-- grade_id: cell-6c05305a2740ebc5, points: 1, grade: true, solution: false -->
```python
# Hidden test for 1c
### BEGIN HIDDEN TESTS
p = Polygon([(0, 1), (2, 3), (4, 5), (6, 7), (2, 9)])
assert p.area() == 18
### END HIDDEN TESTS
```

<!-- grade_id: cell-ba7f2b5975797431, points: 1, grade: true, solution: false -->
```python
# Hidden test for 1c
### BEGIN HIDDEN TESTS
c1 = Circle(-1, 1, 4.7)
c2 = Circle(0, 3, 3.14)
g = Group()
g.add(c1)
g.add(c2)
assert round(g.area(), 4) == 100.3726
### END HIDDEN TESTS
```

<!-- grade_id: cell-7ab5c37b49545f0b, points: 1, grade: true, solution: false -->
```python
# Hidden test for 1c
### BEGIN HIDDEN TESTS
c1 = Circle(-1, 1, 4.7)
c2 = Circle(0, 3, 3.14)
p1 = Polygon(([(-7, -3), (4, 3), (4, 5), (6, 7), (2, 9)]))
g = Group()
g.add(c1)
g.add(c2)
g.add(p1)
assert round(g.area(), 4) == 147.3726
### END HIDDEN TESTS
```

#### Question 2: Defining Vectors (7 points)
Define class `Vector` for n-dimensional vectors as follows:
- `Vector(l)`: Creates a new vector with dimension `len(l)` from list `l` of numbers; raises `TypeError` if `l` is not a list or not all of its elements are of type `int` or `float`.
- `v.dim()`: Returns the dimension (length) of the vector.
- `v. __getitem__(i)`: Returns the `i`-th component of the vector, **where components are indexed starting from `1`**; raises `IndexError` if `i` is less than `1` or greater than the dimension of the vector.
- `v. __setitem__(i, x)`: Sets the `i`-th component of vector `v` to `x`, where components are indexed starting from `1`; raises `IndexError` if `i` is less than `1` or greater than the dimension of the vector.
- `v.__str__()`: Returns a string with a readable representation of the vector, see the example below.
- `v. __add__(other)`: Returns a new vector that is the component-wise sum of `v` and other; raises `ValueError` if `other` is not of type `Vector` or if `other` is of a different dimension.
- `v.__mul__(other)`: If `other` is of type `int` or `float`, returns a new vector resulting from the scalar multiplication of `v` with `other`, , i.e. with each component of `v` multiplied by scalar. If `other` is of type `Vector`, returns the dot product of `v` and `other`, which is the sum of the products of the corresponding components; raises `ValueError` if `other` is of different dimension in this case. If the type of `other` is none of `Vector`, `int`, `float`, raises `AssertionError`.  
- `v.__rmul__(other)`: Defined exactly like `v.__mul__(other)`

Python uses following equivalent notations:

    v[i]       =  v.__getitem__(i)
    v[i] = x   =  v.__setitem__(i, x)
    str(v)     =  v.__str__()
    v + other  =  v.__add__(other)
    v * other  =  v.__mul__(other)
    other * v  =  v.__rmul__(other)   if other.__mul__(v) is not defined 

<!-- grade_id: cell-e2595c6e1f70d0d7, points: 0, grade: false, solution: true -->
```python
# Define your Vector class here

### BEGIN SOLUTION
class Vector:
    def __init__(self, v):
        # replace with your body
        
        if type(v) != list:
            raise TypeError
            
        for e in v:
            if type(e) not in (int, float):
                raise TypeError

        self.v = v
    def dim(self):
        return len(self.v)
    
    def __getitem__(self, i):
        if 1 <= i <= len(self.v):
            return self.v[i - 1]
        else: raise IndexError
            
    def __setitem__(self, i, x):
        
        if 1 <= i <= len(self.v):
            self.v[i - 1] = x
        else: raise IndexError
            
    def __str__(self):
        return 'Vector: ' + str(self.v)
    
    def __add__(self, other):
        if type(other) == Vector and self.dim() == other.dim():
            return Vector([self.v[i] + other.v[i] for i in range(self.dim())])
        else:
            raise ValueError
            
    def __mul__(self, other): # new
        if type(other) not in (int, float, Vector):
            raise AssertionError
            
        if type(other) == Vector: # return dot product
            assert self.dim() == other.dim()
            return sum(self.v[i] * other.v[i] for i in range(self.dim()))
        else: # return scalar multiplication
            return Vector([self.v[i] * other for i in range(self.dim())])
    def __rmul__(self, other):
        return self.__mul__(other)
###END SOLUTION    
```

For example:

```python
v1 = Vector([2, 3, 4])
v1[1] == 2
```

```python
v1.dim() == 3
```

```python
v2 = Vector([1, 2, 2])
v2[3] = 3
str(v2) == 'Vector: [1, 2, 3]'
```

```python
str(v1 + v2) == 'Vector: [3, 5, 7]'
```

```python
str(v2 * 2) == 'Vector: [2, 4, 6]'
```

```python
str(2 * (v1 + v2)) == 'Vector: [6, 10, 14]'
```

Following raise exceptions:
(remember, these are supposed to be *mathematical* vectors so we are starting indexing from 1, not 0.

```python
try:
    v1[0] # raises IndexError
except IndexError:
    print("Successful test case.")
```

```python
try:
    v2 + [9] # raises ValueError
except ValueError:
    print("Successful test case.")
```

```python
try:
    v3 = Vector([1, 'cat', 3]) # Raises TypeError
except TypeError:
    print("Successful test case.")
```

<!-- grade_id: cell-e98b86536ff53624, points: 1, grade: true, solution: false -->
```python
# Hidden test for Vector
### BEGIN HIDDEN TESTS
v = Vector([1,2,5])
assert v[3] == 5
### END HIDDEN TESTS
```

<!-- grade_id: cell-a7272504f3a9ac24, points: 1, grade: true, solution: false -->
```python
# Hidden test for Vector
### BEGIN HIDDEN TESTS
v = Vector([1,2,5])
assert 3*v[2] == 6
### END HIDDEN TESTS
```

<!-- grade_id: cell-547277d9f0f5b33d, points: 1, grade: true, solution: false -->
```python
# Hidden test for Vector
### BEGIN HIDDEN TESTS
v = Vector([1,2,5,0])
v[4] = 5
assert v.dim() == 4 and v[4] == 5
### END HIDDEN TESTS
```

<!-- grade_id: cell-67da8a068242283d, points: 1, grade: true, solution: false -->
```python
# Hidden test for Vector
### BEGIN HIDDEN TESTS
v1 = Vector([1,2,5])
v2 = Vector([3,7,9])
assert v1 * (v1 + v2) == 92
### END HIDDEN TESTS
```

<!-- grade_id: cell-b87295c960c2aee7, points: 1, grade: true, solution: false -->
```python
# Hidden test for Vector
### BEGIN HIDDEN TESTS
try:
    v = Vector([1,'cat'])
    assert False
except TypeError:
    pass
except:
    assert False

try:
    v = Vector([1,2,5])
    v[0]
    assert False
except IndexError:
    assert True
except:
    assert False
    
### END HIDDEN TESTS
```

<!-- grade_id: cell-460fbcafc8b42236, points: 1, grade: true, solution: false -->
```python
# Hidden test for Vector
### BEGIN HIDDEN TESTS
v = Vector([1,2,3])
try:
    v * 'c'
    assert False
except AssertionError:
    pass
except:
    assert False

try:
    v + 3
except ValueError:
    assert True
except:
    assert False

### END HIDDEN TESTS
```

<!-- grade_id: cell-7ae3af0dc7dcc80e, points: 1, grade: true, solution: false -->
```python
# Hidden test for Vector
### BEGIN HIDDEN TESTS
v = Vector([1,2,3,4,8,0])
assert str(v) == 'Vector: [1, 2, 3, 4, 8, 0]'

try:
    v[12] = 3
except IndexError:
    assert True
except:
    assert False
### END HIDDEN TESTS
```

#### Question 3:  Robust Change (5 points)
Complete following function according to the specification in the documentation string. You will need to create your own custom exceptions for these functions.

<!-- grade_id: cell-3831d37862b52a19, points: 0, grade: false, solution: true -->
```python
# Write your answers here. Define ChangeParameterError and ChangeRemainderError in this code cell as well.

def change1(a, d):
    """
    Computes the change of amount a given denominations d. 
    Parameter a must be of type int, d must be of type list of int, and the
    elements of d must be in ascending order, otherwise ChangeParameterError is raised. 
    The result is a dictionary with keys from d, mapping to values of 
    the number of coins / bills of each denomination. 
    
    This is computed by first taking the maximal  number of coins / bill of the 
    highest denomination, then the next highest, etc.
    If no exact change is possible, ChangeRemainderError is raised.
    """
    return    
    
### BEGIN SOLUTION
class ChangeRemainderError(Exception):
    pass

class ChangeParameterError(Exception):
    pass

def change(a, d):
   
    if type(a) != int or type(d) != list or len(d) == 0 or \
        not all(type(d[i]) == int for i in range(len(d))) or \
        not all(d[i-1] < d[i] for i in range(1, len(d))):
        raise ChangeParameterError
    i, r, c = len(d), a, len(d) * [0]
    
    out = {i:0 for i in d}
    while i > 0:
        i = i - 1
        out[d[i]], r = divmod(r, d[i])
    if r != 0: raise ChangeRemainderError
    return out
### END SOLUTION
```

```python
change(3, [3, 7]) ==  {3: 1, 7: 0}
```

```python
change(15, [1, 3, 7]) == {1: 1, 3: 0, 7: 2}
```

```python
try:
    change(True, [4]) # raises ChangeParameterError
except ChangeParameterError:
    print("Successful test case")
```

```python
try:
    change(3, 7) # raises ChangeParameterError
except ChangeParameterError:
    print("Successful test case")
```

```python
try:
    change(3, [7, True]) # raises ChangeParameterError
except ChangeParameterError:
    print("Successful test case")
```

```python
try:
    change(3, [7, 3]) # raises ChangeParameterError
except ChangeParameterError:
    print("Successful test case")
```

```python
try:
    change(3, [2, 7]) # raises ChangeRemainderError
except ChangeRemainderError:
    print("Successful test case")
```

<!-- grade_id: cell-16ccf86c7354d6ab, points: 1, grade: true, solution: false -->
```python
# Hidden test for Question 3
### BEGIN HIDDEN TESTS
assert change(16, [1, 4, 9]) == {1: 3, 4: 1, 9: 1}
### END HIDDEN TESTS
```

<!-- grade_id: cell-f7cd91f11e24e067, points: 1, grade: true, solution: false -->
```python
# Hidden test for Question 3
### BEGIN HIDDEN TESTS
assert change(11,[3,7,11]) == {3: 0, 7: 0, 11: 1}
### END HIDDEN TESTS
```

<!-- grade_id: cell-b4e9cde2bec380e9, points: 1, grade: true, solution: false -->
```python
# Hidden test for Question 3
### BEGIN HIDDEN TESTS
try:
    change(3,[9,0])
    assert False
except ChangeParameterError:
    assert True
except:
    assert False
### END HIDDEN TESTS
```

<!-- grade_id: cell-e1490314fa409f6e, points: 1, grade: true, solution: false -->
```python
# Hidden test for Question 3
### BEGIN HIDDEN TESTS
try:
    change(5,[2])
    assert False
except ChangeRemainderError:
    assert True
except:
    assert False
### END HIDDEN TESTS
```

<!-- grade_id: cell-65653f27801c6400, points: 1, grade: true, solution: false -->
```python
# Hidden test for Question 3
### BEGIN HIDDEN TESTS
try:
    change(5,'1')
    assert False
except ChangeParameterError:
    pass
except:
    assert False
    
try:
    change(5,[3,4])
    assert False
except ChangeRemainderError:
    assert True
except:
    assert False
### END HIDDEN TESTS
```