## COMP SCI 1MD3 Introduction to Programming, Winter 2018
### Douglas Stebila (Instructor), Joey Legere, Karl Knopf, Natalie Chin, Victor Chen (TAs)
### Lab 5 Assigned Saturday, Feb 10, Due Friday Feb 16, 5pm
### Maximum grade: 20 / 16

The purpose of this lab is to:
* Be comfortable calling functions from within other functions
* Practice implementing numerical methods
* Save us from any future embarassment with aliens

#### Practice Question 1: Binary and Decimal Conversions

Write two functions, `dec2bin(x)` and `bin2dec(x)` that convert a positive integer from decimal to binary or binary to decimal respectively.  For `dec2bin`, the input `x` is an int, and the output should be a string.  For `bin2dec`, the input `x` is a string, and the output should be an int.

The iterative approach is easiest.  If you are feeling adventurous, you could try to come up with a recursive solution.

<!-- grade_id: cell-9b4bf2416a059320, points: 0, grade: true, solution: true -->
```python
# Your code here.
### BEGIN SOLUTION 
def dec2bin(x):
    
    ## Iterative solution
    b = ''
    while x > 0:
        remainder = x % 2
        x = x // 2
        b = str(remainder) + b
    return b

    ## Recursive solution
    if x == 0: return ""
    else: return dec2bin(x//2) + str(x%2)


def bin2dec(x):

    ## Iterative solution
    d = 0
    for i in range(len(x)):
        if x[i] == '1': d += 2**(len(x) - i - 1)
    return d
    
    ## Recursive solution
    if len(x) == 0: return 0
    elif len(x) == 1: return 1 if x[-1] == '1' else '0'
    else: return 2 * bin2dec(x[:-1]) + (1 if x[-1] == '1' else 0)
### END SOLUTION        
```

```python
dec2bin(0) == ''
```

```python
dec2bin(4) == '100'
```

```python
dec2bin(5) == '101'
```

```python
bin2dec('') == 0
```

```python
bin2dec('100') == 4
```

```python
bin2dec('101') == 5
```

#### Practice Question 2: Passing a Function as a Variable

Write a function, `stats(f, l)`, that returns what you would get if you call `f(l)`. That is, `f` is a function that computes a given statistic on a set of data, and `l` is a list of data points.

To test, use the functions `bound`, `span`, and `mean` from the previous lab.

<!-- grade_id: cell-456b16f6a5237baa, points: 0, grade: true, solution: true -->
```python
# Write your answer here.
### BEGIN SOLUTION
def bound(l):
    
    return [min(l), max(l)]

def span(l):
    
    return max(l) - min(l)

def mean(l):
    
    return sum(l) / len(l)

def stats(f, l):
    
    return f(l)
### END SOLUTION
```

```python
DATA1 = [1, 2, 3, 4, 5, 6]
DATA2 = [i for i in range(100) if i % 3 == 0]
```

```python
stats(span, DATA1) == 5
```

```python
stats(mean, DATA1) == 3.5
```

```python
stats(bound, DATA2) == [0, 99]
```

#### Question 1: Decimal/Hexadecimal Conversions (5 points)

Write two functions: one that converts from hexadecimal (base 16) to decimal, `hex2dec`, and one that converts from decimal to hexadecimal, `dec2hex`. Output your hexadecimal values as a string, for instance, `31` in decimal would be the string "1F" in hexadecimal. Your output is case sensitive

<!-- grade_id: cell-ffec72cc583c9c6c, points: 0, grade: false, solution: true -->
```python
# Write your answer here.

###BEGIN SOLUTION

def dec2hex(x):
    
    return str(hex(x))[2:].upper()

def hex2dec(x):
    
    return int(x, 16)

### END SOLUTION
```

```python
dec2hex(31) == "1F"
```

```python
WASTE_OF_WATER = "DECAF"
dec2hex(912559) == WASTE_OF_WATER
```

<img src="http://brain.mcmaster.ca/joey/img/decaf.jpg", width=300>

```python
hex2dec('2A') == 42
```

```python
hex2dec('C0FFEE') == 12648430
```

<!-- grade_id: cell-31c9848c88d2fc07, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert dec2hex(1027565) == "FADED"
### END HIDDEN TESTS
```

<!-- grade_id: cell-048a535063da5366, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert dec2hex(1337) == "539"
### END HIDDEN TESTS
```

<!-- grade_id: cell-9eb0a498ba76d567, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert hex2dec("BABB1E") == 12237598
### END HIDDEN TESTS
```

<!-- grade_id: cell-74672c8fd795b3ce, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert hex2dec("DECADE") == 14600926
### END HIDDEN TESTS
```

<!-- grade_id: cell-f549c11cc3024525, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert hex2dec("BA5EBA11") == 3126770193
### END HIDDEN TESTS
```

#### Question 2: Fibonacci Sum (5 points)

The $n^{th}$ term of the Fibonacci sequence is generated by adding the previous two terms. The first few terms are:
$$ 1, 1, 2, 3, 5, 8, 13, 21, 34, 55 ... $$

Write a function fibsum(N) that returns the sum of all even valued fibonacci terms that are less than N.

For instance, for N = 40, we would return 44 (2 + 8 + 34)

<!-- grade_id: cell-bc772084dd8eaf72, points: 0, grade: false, solution: true -->
```python
# Write your answer here.

### BEGIN SOLUTION

def fibsum(N):
    
    i = 0
    s = 0
    TERMS = [1, 1, 2]
    
    while TERMS[i%3] < N:
        
        if TERMS[i%3] % 2 == 0:
            s += TERMS[i%3]
        
        TERMS[i%3] = TERMS[(i+1)%3] + TERMS[(i+2)%3]

        i += 1
        
    return(s)

### END SOLUTION


```

```python
fibsum(-1) == 0
```

```python
fibsum(40) == 44
```

```python
fibsum(4*10**6) == 4613732
```

If these cases run too slowly, you may need to modify your algorithm:

```python
fibsum(10**1000) == 5933608372629145798383544242983334636899291050047879463824293309987965343882047512984107588698285346632851981219062849855970529781272597133037980905941846567381108185609155598002212061744588022560666944282767462121189302686560263335164922661315868839451963485338930580620175723568033024082499799721271328257452544308488139652872804895873257816488895097469482618389027664983663019272178104872928427579529466738208129384632199186931292053505993390945828326147177151692121336204311895165981982728598087114287157410488507274530820653725550887083368470109297084168625855256569091543118913762196588623005900476707497335157598848209727884494346486850096686339118011583322943230155678188177779582642187147830838023871251508179354174068722627132322379667374013645021983195445921872203922884810130060459330632124749284199708376404669104869936023808844711242768526994447908900991933324168339513635421902151293084025917812258411608177117040739665776652404631304245925539202140227103643288849790111066129620913716
```

```python
fibsum(10**10000) == 6699559432585660545573702667084889471167548773597601986481641090446197952712804461154129282458711835317064458448149169234470011778174418365079734961165609520696228859092291252739958860174970425289913896373982524025730492299535004427028788703779885176083917699562860918497323462333709760892065259929895842180757940776104907131526748923152364416432274806937079977677152203430758371528777273426810261967496544308254986764531933263242154357764546191415817871614846861410010482568300686033478450295813207065383930067909263776045829506906129092493811782915469448125217934476410108381371560131575694900421937937560793969157346662004140263746686287943354704134224487565966502139965935879931203379811054522254417930205809519962200030648267537578828887461807677288341217928560065635529809314994693648096547150895428401563383973912241282496393193068676578945003920255142257645513082935697624070043084104895184744821027473980767787959157719864153354417489834278474794668828180652816569839652583760076563703007078472828572998387858294441265495505964168666491283334693208848992018518221671107638109228908869463915017151926499700223213224746894993224139274484035253769394889358025482734683133711789172092180687346523953211299462230355746380413684081815958336520812595742744354364850026413643515012649060802696371835896948898750049841914010544083083072143246423681512607160467315894898295393233731816803814595968050847575109047815122974769489219193144703390472916521006261010907793055236376575453887459395639693961814495598526618907113476832123887477894382557666012992282179871291158403339047578096518489586872398298361528187953935855765682840184864556339364751264452546895815678757400830400758456954336571114383511030430430794865619730836206946139364700889555625225865142184498975714617870760212990468576194430907636620364096398439597120565745988643079868609503725477137311957820103190906982178448087223104484982565092982648391709646601601791307302935277793631741248483769465307112681190288665573979756550097795102428679915979257134233678903459868627459525914768637313194856921878336607599740938923062677480658374923951805349002395236508974135704462472175524478742617796572732355812854424929969280235441641960343543904339192590624422843681310080332342226809819878986830041672298016768052225250348318407356186347677108763668998431288567688009001266741322006532365900296067682260130081568414862371875664103738999873281260627320325706242390648611413318784206249461336653396937042196775055292480866866591115960083117497193040690873559862170589785978696196473142297391842596152633213093903605610548182448333197167038522718591353628681348214454673350189865898820618864508773913996348110939986838538689839328877757567821931467012412537971231528991193848334008880351430281432818790343536478181741983038785690626138319924647108106458281154569062500163473398347614975054135060149377334043994118223722246442606760221866570623359717528463564705561853439714275146144969584625020943157778215630375060089096749295342049338381533002976126537378944413398672066375834861835762057157099264536572103783479904442950533478701826005029803841154401620961421125276962933914155673435205853673583907003757383309908493424419862295711566563838234823842722319407517053871452845280379708913288894561078584464866883826792085163661586053569401851953413998228213415637049526138524879393807446983305109251852283153180624526155345296659236656891901436548117351102579599725494281283929955912695711428127526284225588804721014729687551779347443862764532533993953623881441901017990643430104597139212120594922897417433193240487708336432182905721763243738145564282991309968962865200024633437178880407621147430625263884239829630337420160193585599777115337471245897099909830144309084446072429210890853956706178522130717899875105591623993227084619686915726303968454846810812670910641357364770546110433601685683480388214666739440248788851836508633069618674226930206401723886483111308210778027706588902645269705537972874848745123910782216388048765150308545688383154995779589249039185247688856002405772637327226486864340724301285425082979895014736052585112783013367875696638479347600539876372518656375494020110587734437524353447554756859696091918896851713170613875456943054995037309977790716348039998043245276462508730894047629490960661782258669711094301970252719679006479334288175177335983574224270993527811768548708772075619493381336092917675136608843900626483139215931735267203655218998357629060030641274537569917230053462033808886826767201687603126084361046489811881582827084919364465480232232251683703107659656867178098892775950447929021918720085882031855068007971910243865242902075209059611623530537225059270750341361623796724713806301467817785132728262662226056779766297865157818714225935738308207501987797603936690939113622796672842950184465387965779696290078976543433942326993874949136719649202440419351467027761851293468408746906932696468309139998639700294794262924439259358189546727906532919134254184960821986317229123959039021651533906261047559196081439365275043809947866893155396778240828938452827009514373770424210738029103780204607226941485380550269553596192027571314080503946672308472895299520943618752539005389025994285345911539975346687685448147689081106537019413221517850836230579544698672360066018132535326689998380550955508916517047750821423678263550998622012246648229375038463105509123128568168022590262248224699681122169817372759524255332271133838606814912160708460548021679967069663290500079680201115907611367397041632063803422122357419262793954275973504580189973928458303181286157282465514438359378712572836027224886460302217813023722344650044057284373037487331974538440444839257007002572942746386190328757889959553302808368593859572496818989771627875831404115942324089299007043590537188304600876430370101246177693864646875293368408075546149587532590146046307940143145170259501879633378727832654396913119090794024193127426722294291899129526635509666401346566324117080504880813531517362126302071547638039707703575512984865879662751568944248849455529859549053781890637080407414480694198012687675408000347019626333805884623594725397397730882037835972361331193369935623856328468216868876766402607413883325782227211118944183083842508312107435194526558237381144780669263870481952249017823111001456682825902993323148491237405846285081096000215143429420743467037656605041955197606401039182731045209160246658828013084890749968152168734329395729088059370351397750525386261464976079891150939692145791760536143844126578900694666962116098870431082102338004405548940986651207500829795972633670482911751493498610179051655782606778613149638713160687859706688923815267934301227748234419085021114600060456895840445895000903707181252856038191097334181412632990017097904730407991059711867580546831124923550333421062940026709886457698947676526515488731123056695099561280429415974907155120930582165077408142538804997392401729340198252443143751840351938500180471396451078043318022813718525130286568397205838434355142809971390203949849831760121543823393349516738784264409497647428783194900046540042072745169037634588637390336441928707267203446713157293984992173732499395884667904178466265342071149119668064699079051401991574802966746313619750996181631638521807567859191537471504450106077904830388223985737292157929614076428959503687851318743103880221469954865863049241452229342365778087324280964775732538556178064700197508263897772866435128369213708932124034085956096768915569504350568736449657751587771610087990855977086120805491270392230637649830079774164498823424403912605875319864787901423755935536894581951433343329145041012051200981283605184142303066515938264632122976606787542649050410299400711993187424144272506244337780042757958484122407927715877279658014111507101395198300204643641886162712088419887930140636880896994303365252882689666771476322707680220380822274866778443158008895697428385492108651753047195976409203756775278673160974987756945520799910746384899971127077677486318968559764302613524072622982596192406661564082048548274661925417569071175247185852144690359753675001778242513570812875852740129530748136198317684653252138703345674732501666256251003788670603938442541650629585053027738553956505574966258724539333669832347350522952097747335697294695655865521656181187694387205959207181537234540197370813746054560033123001849805594961388024724565960983729629002358765565333611303244381693649142415269848935868243717355346336460409513530433094860336611705776830939333438532297716119202154301990437527773135049678495979764589356011715483187962185328174222959256195971428170815356408466016826581970396276428724172892576788184654653979819444788229970870741912298827873406076383194738696590669188154115423719759725203777926023297448041758772516848642379681098529769752380741468104663578064864799971534692534443227685295164302719841786337051963656768742393378748012422637743430288027634630288072943822608605929100738364315388908364884835836327577964617179359436199853428852741942587346967906779362664809098615023056549149794913463129040849583238476429861381692803234207508556025615112734905492522439297607498306713034372157913178165979795242898485438164077736573803987665453800099056067140828404012478713596309055717476071052120275136594046843455687207699703111604117749910097409163448796955504824896863000639873956766352420970842700580538202873290424515056567139531897749411262738122903567614462001583826111124877203237313838208019106439757803145527852815064631536931485283281359975039172147234848145136936155183824227415239485723468598969988055850503488605463345369990708368814762044947036329786113964138824466328406620819683485006156058462937935225514534159537697116657426746747000390032900809734420481630696958310951513291542821900228710825434801996384090603164466117688682144730535006939498525816301981771991943271084985657478229818815217763847027854263296727301956426372767666815542776283790115944073599765938303798284803712754269377654272283866744578623692932087544881833217085767633512288979009729489670770865254341868242923648006640258380974
```

<!-- grade_id: cell-3b30499e72fbe01b, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert fibsum(0) == 0
### END HIDDEN TESTS
```

<!-- grade_id: cell-db1809540893aa53, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert fibsum(7500000) == 4613732
### END HIDDEN TESTS
```

<!-- grade_id: cell-ef338cc40aa0c02f, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert fibsum(123456789) == 82790070
### END HIDDEN TESTS
```

<!-- grade_id: cell-e3cd81e5a6885828, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert fibsum(10**1024) == 3961957999672379439840769192827030105457313390675612001248064887482835737210631587414596404349120013232818481232897807416011058430907788668949992747822531697100611769717312384710513049044104338410786436827045995781311911378330887822706273602906349721450136688623608658567262956481837774314320695276172283671687987488599027549009710985135218538904419776886206736263289018654705256927755612283958742321229235447436960269858111092129075522063525524139964492705339321422189025432643961048973937851702046819016255630293581023579189191864932779876672014747838250176752619096035633006826878239821515434232834545810525168025226088033916269367773909423038108921769148739541582724644085798059315650255894453981487582178407518494572287566926608929700459065946359010802125820161535870484026071677669928420811832513093461995180003057602558834479852093395967994073057948628120376456224641705927904258069050400463241329327814233713491125457710542993968091518000109066567256492619259020792967017323295529441206172968283349961762575442460450
### END HIDDEN TESTS
```

<!-- grade_id: cell-c9594a91e76e38d7, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert fibsum(11**12345) == 4386824498096412299864976180374989092963630116217555841935388302263153643505281150691199849799914644322530241671579176951834644263875272054035899135857000910927255108451530204624843157906039544628249183546926618564544294311384172219726257199421050948823507580219746166008281213889538485523181889355805531244060466230685285427931925253856064642495284591802961705763182566370177207157234275081580707156508199815261896521001101199892509941875029727616841431242137329943764877746257740351355642043432141199283645010178479786164634580183432612624612238597708571974314846214168766977714074695061533291269155230906380959359364914600610473545232556003003389803536995264414932512556101670215173045873021844137829515163917322782030353364098109644420847497917643082231480991159790140790199620878930462043345547520024765447324985256200806102028877060917969875135633792544602544427524028313402764003271976603990655864081132583501748353283203561978046935399868622830014608537318834499658462310476086447366734782433231649495096197884981266240675104082532675860275055231106639040236509524004403356057621675413773237505497171079965990807581261840592733713166137767319866234124201344331669506483218878888432248714072967742292558647322351113358199194955338257970731913629728331527480453849784348215193333573406914085614106756374029437587266356615527558854173265930039594331084527447990158531933911995871678423798921574503953656176176709259685109855543836849570536961123675546622582503592489749250561059010583751877040137088744596015978161010539515562304540301859339176014448036316596863781711297484661021536156296857585821411415070842636991896380407514762903605251453851960058819058618570353316568082359212099426359389364576955555857516267251782749965866732976793692736604535492655081232949094399343736229604373664423056369247042603546308991421335377954724654755953175048048701295208491118741665981721392143838234865339254317034170969451935143205060443347260042841368376068601479626224324801433375305056324638461744127482234287278996231234080338200497140052098290425670112042999861158900721642397639348846718882608766812046548290141218185216968411275191795801971503146915918261260915074720671427543262046717297123143778066634468125162208970704339006066473876545193068929472015086705330653646996084191867871909890896405552621522211861595747278581169742960418530531529975752430594178065700632105693200320366809992149842915240424902142002236192493295189925459093062129719765716021456802587439875408466136312575204330008265520208383478251941813816870533774503491269807263066127824044343238572589691458395876988139703129356245876646693442673892259324766452020034496292218164965782430595807778002292446634732480648409752235049964711257530238661681923218431260907469398762579978463122159463538771663426186678423493717954226437489222406677627143661792158052081744472694465516004630079598506183958856474956110188142281005755817327457969426323098233579689877903991318016954682296086010929332128312474121677370838255435731546030368513302975323358217449764539734875171576238286611851776966655372253515682289237513927435819471347721311411811113603185855567501478605521624087623342134026509479293754462698656871036759177345690209181580339670177590892241430273795635915188584697341592018271490071970926725920612591963737611989804872170484504301575625591766567665417349430147523936755344102043286823092155344094226961453497207820254597748090299943506171486138902591162552607564125118049650055723469520307819457383120896810352951941392339050841869921101080353670860217873150130582759385529895869546615535029909818242180421812277228242668503323354565176837304408808046575968483285336280711409777900615939839050551081165959917562466001966592945050454008019126102710833221770120801709382764458136610222471925800139197978660116578874880909098315145543513701543443930379927430346317121973790694282060998598261567189698312975619532071301130868303136749759822298472568628159565049225490993081317378895942169967016644182866940631343596692994744107835625228798429320333312819522866670977451767206598344665806658391290905869976103870457982862681102835353466667639723374994169024773234623201991782145978788625083437354918279456351757773182112455853180046971898617533300815530983681044757706060755544870888063494292628482571637203896388748060574308568707747649764073579882060360820698124273904777725682803095011762184265973123360657941268326379021979790883884692036910816732345118160513184933195257620939460682149515907679332032398778120368489351010453747231354338929248890253602203854374572091032441853526953735886559006992309183715737327712846675582294891819410736792430928162891776278030090521180284085470417065255561404122137350993470507389089442249424099818224426073589148532904705785327582269100275048192171488738414582443576111273823778204991657713714773868785775973598896726651493931671020052615283966496139877897660892042901907413938078346642729090579605098168594980017518817863801187364856158272612996426999371658045820448296147238238058771498616938022685861239467487659428970196813967271102265253626182003555350305725542772761520250953320784899466527835648099872626287336815504835490390772315855764662367178627892179142670706496231576160172766778165308068549750257128878640949812940806067628780653917478834272881185886744196326689784007330941133159798611476311463092284835842124919418160194291673924504244837487956974704195153568072464191743371825244002434139563227413815466057445269017154056806404741601571051034864348798251521619668183867423744565208878286083316277355357316781553784040547998257242019448134941768553946284164852127954160167121340519294791109196810836763537260383753125983596207393989644861347632437891961043725930649619979691101647129659229902880808189434835116149178439769814240881577491437868374138031260401333564018107538830705248177668126715210023214242422999742202747396484891742676111300611145335658961060461079841410253391883878987916174539980448746722674187048577401610649342851275377759215902283507097618151086009608923769228670418821939399802586120499114694344146275198357698359135048916082747563020119156110186802412678083604023014225497524682561329665001807008052574759795051876956489741377487163824886902780918453198067096930666144327531877286778601559955504067602277384221917959396525004764384471217200182287054743392436431279182770145973280376091046119323048057293033800027445871728542992891706751114204470382067728935107487843691440794545635476519713323354223796792194628063454916644930983118202434561389268624980571136173917768488221566417373315535021861840544943532289169934804522069279010943389389576975067581358525014796271774704986076085821139010774055280483520465079986562578457229955553109936914268919545516244049270649268045409515710211095060540534682789787594334613883493934203596914025846195663283868628321141933607035499182956923841713284024832294909426557502546392012186862021135550079472400857648172438027692057800782362203394602487370639004497125434961329857341928179728435501092683931246292751904918345837157151996678755883832572069541544232811133268952851306130442803576009943876184746320744883631010764199064351360712306336137926683553749618969009460922719393257102710231063168105036859226136963509997889227257223024214027373304299359408388374761896440009044071793236927373381188369237691751038749968405745068209498919188316700179850566463183982828178769636972458995767565958205189822466911064568046178920768493846941956237173946718374604689891894198297092087150374104758138009655895243645766046638575150592750369096338180593489268610529562482298820492092569024340255395507831281843013408008029193727353523905396496220921588751920494034099653550294669416794431197061227180000133371572796691573979083390876347147442817051137314779423949339876195438000195756912917257203974104274910014175613373988527448534593206621374906703872927136332260595251054586920106231967646050959467299208215355078473311704115599307798355598475179173286414997110525114157111158754786247395175190676802590756865371942822110461723630723567927606773454246599726138028361986358299444177314540060298182689796330407710392381188687407071254610662752593088628815009544383148929830742457864453798417598698356385587273165259976269508762175719500319104130705149056399651026660079751575937519428822458706160043824814231853073637642239849486357898105068535256226612877653995756386004111302237494690495148020928748638075569673929317589754761091907327877137580286411139900538114799156068633712124555202767520498485058823388863872070087818702381706418945077170326446618046616410144594705953664665019930155115179365346800902616871271640910789399365011758464775193854301132208248460878071411352051661835106245742788695614278830560419879723938432011629862858441624441658932469689292079880811982194377305860230970344298604954986541439266634647588820050912909944453270635161001536881835350935245636035157010642643393677254710664782418126386598585203998971979132282429429879183269181533835286348152285228148505993929136286281553948590646974847740421801477746716886912529336484383622398927993216311068197413316687137938006931141916033490121165282627127440413546386566018040942497378136203279854446785509432945018240117047685288764745211831689400519360939380439000619954449114961667774978281474935797622122213929092275376594220689408550086831143268631225650035567783144564601796649229283718507718586416242430027067428637893123316593372899416665352221526091535722963883217956700648316810644548935649154932663012269345139417382009641126304992580359238530143937644136924382387290399991639718862929581430566441752923019926765347355184927207679400403075502042737662598798079013022004167707633155521108475951053495660956298701497441779724826870630097266016323647849388287244384688380939998894173479481457612223956896723251511840540062813364654261172493822277665564612120056036407397693836803791734809573183133643009131055650763389849216715622933881416705253011449825258233776940416056992471954550044877719312412259162953612653745353658481947612768189092630704695155546125119314809420810903647061555560552364692293383207661009023955303550691714977890620720447103430138098670424758537273980401466623112810510514243324500462369801560867920662072574348815201855034379932691076278640105633577816932967367221488242465347419180338240473154075580517110159377302349664319294055315038475491595960254580869625175617322078451493661476136509955308322031077021819192857933502924666188132441494768408905089498524376410240877028267871625109745339441025984622838209210963365879532323868573852713246478689564112420602931462787392364701252312705950829008081113326594879908529430825302543125294507078809568957363131522739729733957938335076990401763547858353240335382410222754537265278465500632495736095869249999044451279951009813457492985187992165689457167759964470468155798355315757370122694744566439652399406767658094431555792769395900991097281213648184792505932235708221848046493090147428799730406441181438682164566063443631813701158562625682042360351546097925853541754945401099036046898497954758731706801239664349542231284544123935971414939790153729339292047005039296186226599772138749226602634890769185365638007808066394968627275087088720383758659242445143042226268418384514502652149267065363509692288916916561177489415054112806012140415122621918041408779524331907787701116932487356597194787922255813484336582794860964652266887109774803529484047203656795415984057807518539574222326186521707955113578172151931138319005411294047353085884862274104562883513950307312171244823807617793111546182612065832856318134892453367050652126210497483863566163558967924807408466470064411830991493538874850357620662636525725892678453856224132113829413835781665084108265445601247669348218960002405730825278434228801915264236263090374212298086669015772749055139374757343402440538425312499325630309715068455553230509120909237874961506257923615637242063518448398194982195022030468826619427183209345143600130104993601672091975106099723759008372673419592620289525424779431472853124150600976690407501060247129159979810894109577764517257675517655493459001549404050121162662478333924740139425277007652721413366830934548754746110087183933251720948364096798553688350924206471632703587302129019988789723240531978237005889073300959580199542206569733659212640153743402414798887777856365162249774235683037759116367798752915260292285729611811876493385660433087056009761314119130501388391431902134068054811426316353971748561106112548314089001897642003824412980199245442347234093624701779367268147133489113303446705327881702579259481154976973179321165043557343774467963941653097937362329423086728498305532968687108903055649570825217946645211427385094785485238640781967420660874436825074385584819866837240003402236230533394434584240962361979615519531888451256924971629246586065301585626857527469952997987810090788407116599151562808691097753189020298
### END HIDDEN TESTS
```

#### Question 3: Newton's Method (3 points)

Newton's method (also known as the Newton-Raphson method) is an iterative way to find the roots of a function--that is, the values of $x$ such that $f(x) = 0$. The idea is this: given a function, $f(x)$, take an initial guess of the roots, $x_0$. You can then find your next guess as: $$x_1 = x_0 - \frac{f(x_0)}{f'(x_0)}$$

Where $f'(x)$ is the derivative $f(x)$. $f(x_1)$ will be closer to $0$ than $f(x_0)$. 

Your original guess does not have to be close to an actual root, it will converge regardless.

For a more detailed description of the algorithm, see: https://en.wikipedia.org/wiki/Newton%27s_method

Write a function, `newton(f, fprime, eps)` that takes a function, its derivative, and a tolerance as output, and returns a value of $x$ that evaluates to approximately $0$ within `eps` (that is, `|f(x)| < eps`).

Note: remember that you can approach 0 from either side. `f(x) = -30` is most certainly less than `eps = 0.001`, but much further from 0 than we'd like.

<!-- grade_id: cell-6e59a5187b224404, points: 0, grade: false, solution: true -->
```python
# Write your answer here.

### BEGIN SOLUTION

def newton(f, fprime, eps, maxiter=1000):
    
    guess = 1
    
    i = 0
    while abs(f(guess)) > eps and i < maxiter:
        
        
        guess = guess - f(guess)/fprime(guess)
        i += 1
    
    if i >= maxiter:
        return (guess, i)
        
    return guess

### END SOLUTION

```

```python
# Defining our functions
f1 = lambda x: x**2 + 3*x + 2
f1prime = lambda x: 2*x + 3

# The test case
abs(f1(newton(f1, f1prime, 0.0001))) < 0.0001
```

Note: The above are lambda functions, and behave just like regular functions when you call them. That is, calling `f1(7)` will return: $$7^2 + 3(7) + 2 = 72$$

```python
f1(7) == 72
```

```python
# Defining functions for second test case
from math import log
g = lambda x: log(x)
gp = lambda x: 1 / x

# Test case
abs(g(newton(g, gp, 0.0001))) < 0.0001
```

```python
# Functions for third test case. 
# Note there isn't actually a 0 for this function, you just need to get "close enough"
from math import e
def h(x):
    return e**x

def hprime(x):
    return e**x

# Test case
h(newton(h, hprime, 0.000001)) < 0.000001
```

```python
# Generates a random polynomial every time you run this cell
from random import randrange, uniform
#from random import randrange, uniform

DEGREE = randrange(5, 10)
COEF = [uniform(0.01,10) for i in range(DEGREE)]

r = lambda x: sum([COEF[i] * x ** (DEGREE-i-1) for i in range(DEGREE)])
rprime = lambda x: sum([COEF[i] * (DEGREE-i-1) * x ** (DEGREE-i-2) for i in range(DEGREE-1)])

poly = ' + '.join(["%.2f*x**%s" % (COEF[i], DEGREE-i-1) for i in range(DEGREE-1)]) + (" + %.2f" % COEF[-1])
diff = ' + '.join(["%.2f*x**%s" % (COEF[i]*(DEGREE-i-1), DEGREE-i-2) for i in range(DEGREE-2)]) + (" + %.2f" % COEF[-2])

print("This polynomial is:")
print(poly)
print("Its derivative is:")
print(diff)

```

Note that because the polynomial for this test case is random, it is possible that it has no zeroes. If this happens this code cell will continue to run forever. None of the graded test cases will have this problem, however.

If you would like to remedy this problem, you could add an optional `maxiter` argument to your function that allows it to terminate after performing `maxiter` iterations.

```python
x = newton(r, rprime, 10**-9)
r(x) < 10**-9
```

<!-- grade_id: cell-ab5bfc6fa8d874ac, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
f = lambda x: (x-5.3123)*(x+12.3456)
fp = lambda x: 2*x + 7.03337
assert abs(f(newton(f, fp, 0.00001))) < 0.00001
### END HIDDEN TESTS
```

<!-- grade_id: cell-1107faa1c8f9fd82, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
from math import sin, cos
assert sin(newton(sin, cos, 10**-9)) < 10**-9
### END HIDDEN TESTS
```

<!-- grade_id: cell-26e40b8d4be45e36, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
from math import exp
f = lambda x: exp(2*x) - x - 6
fp = lambda x: 2*exp(2*x) - 1 
assert abs(f(newton(f, fp, 0.00001))) < 0.00001
### END HIDDEN TESTS
```

#### Question 4: Intersection of Two Functions (3 points)

Write a function, `intersect(f, g, a, b, eps)` that takes in two functions, `f` and `g`, and finds the point of intersection between `a` and `b` with a tolerance of `f`. In other words, find the point `c` such that:
$$ |f(c) - g(c)| < eps$$
and
$$ a \leq c \leq b $$

Use a modified version of <a href=https://en.wikipedia.org/wiki/Bisection_method#Algorithm>the bisection method</a> to find `c`. You will have to modify the algorithm such that rather than finding when `f(c) = 0`, you find when `f(c) - g(c) = 0`.

<!-- grade_id: cell-9209a400fa396373, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
### BEGIN SOLUTION

def intersect(f, g, a, b, eps):
    
    c = (a + b) / 2
    
    while abs(f(c) - g(c)) > eps:
        
        c = (a + b) / 2
        # compare signs
        if (f(c) - g(c)) * (f(a) - g(a)) > 0:
            a = c
        else:
            b = c
            
    return c     

### END SOLUTION
```

```python
f1 = lambda x: 2*x + 3
g1 = lambda x: 3*x - 4

# point of intersection
x = intersect(f1,g1, -10, 10, 0.001)

# test case
abs(f1(x) - g1(x)) < 0.001
```

```python
f2 = lambda x: x**2 + 3*x + 9
g2 = lambda x: -2*x**2 + 10*x + 8

x1 = intersect(f2, g2, 0, 1, 10**-9)
x2 = intersect(f2, g2, 1, 3, 10**-9)

# Testing both points of intersection separately
if round(x1, 7) == 0.1528729: 
    print("Passed for x1")
else:
    print("Failed for x1")
    
if round(x2, 7) == 2.1804604:
    print("Passed for x2")
else:
    print("Failed for x2")
```

<!-- grade_id: cell-b97f26dd1961c7c2, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
from math import cos, tan
z = intersect(cos, tan, 0, 3.15, 10**-9)
assert abs(cos(z) - tan(z)) < 10**-7
### END HIDDEN TESTS
```

<!-- grade_id: cell-d9c21b0e6bc54327, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
from math import sin, asin
z = intersect(sin,asin , 0, 1, 10**-9)
assert abs(sin(z) - asin(z)) < 10**-7
### END HIDDEN TESTS
```

<!-- grade_id: cell-7b395fcf8e40ede3, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
from math import exp, acos
z = intersect(exp, acos, 0, 1, 10**-9)
assert abs(exp(z) - acos(z)) < 10**-7
### END HIDDEN TESTS
```

#### Bonus: Alien Palindromes (4 points)

It's an exciting time to be alive. Aliens have made first contact with humans--or so they thought. Unfortunately, their first encounter with our species was with "Starman", the most recent endeavour in Elon Musk's desperate attempt to return home to Mars. Needless to say, they failed all attempts at communication with the inanimate spacesuit. Naturally, Ol' Musky was quite disappointed.

Top scientists have hypothesized that our failure to communicate with the aliens was because, well, they simply had too many fingers. On earth, we use a base 10 number system because humans have, on average, 10 fingers. In order to display our intelligence in any future encounters with extraterrestrial life, we want to show them that we respect however many fingers they have, and the number system they use as a consequence.

Your task is to write a program that finds numbers that would be considered 'interesting' in both of our number systems--namely, palindromes. For instance, 1991 is 7C7 in hexadecimal, and both numbers are palindromes. Of course, we don't know how many fingers the aliens will have, but Elon assures us that it will be no more than 36. This is convenient because we can just use `0123456789ABCDEFGHIJKLMNOPQRSTUVWXYZ` to represent our numbers.

Write a function, `alien(b, N)` that finds all numbers up to N that are palindromes in both base 10 and base b. Exclude any one digit numbers in either base--obviously, 1 is 1 in both decimal and any other number system. Providing the aliens with such trivial solutions will most likely anger them, and result in the complete annihilation of humanity.

Your function should return all palindromic numbers in both decimal and the alien system as a list of tuples: [ (decimal1, alien1), (decimal2, alien2)... etc] in increasing order.

<!-- grade_id: cell-77063b4697eac250, points: 0, grade: false, solution: true -->
```python
# Write your answer here.
### BEGIN SOLUTION
ALPHABET = '0123456789ABCDEFGHIJKLMNOPQRSTUVWXYZ'

def arbitrary_base(base, x):
    
    if x == 0: return '0'
    
    c = ''
    power = 0
    while x > 0:
        
        remainder = x % base
        x = x // base
        
        c = ALPHABET[remainder] + c
        
        power += 1
        
    return c

def is_palindrome(s): return s == s[::-1]


def alien(base, N):
    
    matches = []
    
    for i in range(10,N+1):
        
        al = arbitrary_base(base, i)
        
        if is_palindrome(str(i)) and is_palindrome(al) and len(al) > 1:
            matches.append((i, al))
    
    return matches


### END SOLUTION
```

```python
alien(16, 1000) == [(353, '161'), (626, '272'), (787, '313'), (979, '3D3')]
```

```python
alien(7, 1000) == [(121, '232'), (171, '333'), (242, '464'), (292, '565')]
```

```python
alien(29, 100000) == [(929, '131'), (1771, '232'), (19191, 'MNM')]
```

<!-- grade_id: cell-e58bfa8edb7ad7d4, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert alien(8, 10**4) == [(121, '171'),
 (292, '444'),
 (333, '515'),
 (373, '565'),
 (414, '636'),
 (585, '1111'),
 (3663, '7117'),
 (8778, '21112')]
### END HIDDEN TESTS
```

<!-- grade_id: cell-7a3553c87274e988, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert alien(17, 2*10**4) == [(252, 'EE'),
 (494, '1C1'),
 (545, '1F1'),
 (767, '2B2'),
 (818, '2E2'),
 (989, '373'),
 (2882, '9G9'),
 (4554, 'FCF')]
### END HIDDEN TESTS
```

<!-- grade_id: cell-afabc9a5ca6248b4, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert alien(13, 2*10**5) == [(222, '141'),
 (313, '1B1'),
 (353, '212'),
 (444, '282'),
 (575, '353'),
 (666, '3C3'),
 (797, '494'),
 (1111, '676'),
 (6776, '3113'),
 (8778, '3CC3'),
 (24542, 'B22B'),
 (25452, 'B77B'),
 (26362, 'BCCB'),
 (56265, '1C7C1')]
### END HIDDEN TESTS
```

<!-- grade_id: cell-b06a6625135faf48, points: 1, grade: true, solution: false -->
```python
# Hidden test cell
### BEGIN HIDDEN TESTS
assert alien(34, 10**6) == [(525, 'FF'),
 (595, 'HH'),
 (2552, '272'),
 (9119, '7U7'),
 (15551, 'DFD'),
 (16061, 'DUD'),
 (21812, 'ITI'),
 (27972, 'O6O'),
 (28482, 'OLO'),
 (31613, 'RBR'),
 (32123, 'RQR'),
 (38283, 'X3X')]
### END HIDDEN TESTS
```