## COMP SCI 1MD3 Introduction to Programming, Winter 2018
### Douglas Stebila (Instructor), Joey Legere, Karl Knopf, Natalie Chin, Victor Chen (TAs)
### Lab 10 Assigned Friday, March 30, Due Friday April 6, 5pm
### Maximum grade: 20 / 16

The purpose of this lab is to:
* Use scikit-learn to implement linear regression
* Implement basic sentiment analysis
* Use machine learning on some real-world data

#### Question 1a: Nicolas Cage: National Treasure or Threat To National Security? [6 points]

It is a well-established fact that, at least from the years 1999-2009, there was a strong positive correlation between the number of films Nicolas Cage appeared in and the number of people who died by falling into a swimming pool. That is, the more movies he was in, the more peopled died by falling into pools (specifically after falling in; the correlation does not hold for people who entered the pool of their own volition).

...No, really. This is real data:
* Number of movies Nicolas Cage appears in: https://m.imdb.com/name/nm0000115/filmotype
* U.S. Mortality data: https://wonder.cdc.gov/ucd-icd10.html

Using the file `cage.csv`, explore the relationship between Nicolas Cage's acting and these (rather tragic) deaths. You may wish to consult http://scikit-learn.org/stable/modules/linear_model.html 

(i) (2 points) Create two plots: One that shows the number of movies Nicolas Cage stars in per year, and one that shows the number of people who drown after falling in a pool per year.

<!-- grade_id: cell-50de3e705f50d954, points: 2, grade: true, solution: true -->
```python
### BEGIN SOLUTION
import csv
import matplotlib.pyplot as plt
%matplotlib inline

years = []
n_films = []
pool_deaths = []
with open('cage.csv', 'r') as fh:
    reader = csv.DictReader(fh)
    for row in reader:
        years.append(int(row['Year']))
        n_films.append(int(row['number_of_nic_cage_movies']))
        pool_deaths.append(int(row['deaths_by_pool']))
    
plt.plot(years, n_films)
plt.xlabel("Year")
plt.ylabel("Nicolas Cage films")
plt.show()

plt.plot(years, pool_deaths)
plt.xlabel("Year")
plt.ylabel("Pool deaths")
plt.show()
### END SOLUTION
```

(ii) (1 point) Use scikit-learn to create two linear models that relate the number of movies he stars in to the number of people who drowned after falling into in a swimming pool. Your first model should use linear regression (also called ordinary least squares), the other should use the "lasso" model.  (You can access the lasso model using using the `sklearn.linear_model.Lasso()` function; the inputs and outputs are the same as the `sklearn.linear_model.LinearRegression()` function we used in class.)

You might need to use the `fix1Darray` function I showed in class for converting a one-dimensional list into the format needed to use in scikit-learn.

```python
def fix1Darray(L):
    return [[x] for x in L]
```

<!-- grade_id: cell-d274bd540fa3d939, points: 1, grade: true, solution: true -->
```python
### BEGIN SOLUTION
from sklearn import linear_model as lm

n_films = fix1Darray(n_films)
reg = lm.LinearRegression()
reg.fit(n_films, pool_deaths)

las = lm.Lasso()
las.fit(n_films, pool_deaths)
### END SOLUTION
```

(iii) (1 point) Now let's use our linear regression models to predict the number of people who might drown based on the number of films Nicolas Cage appears in.  Then, create a scatter plot.  The x axis will be the number of movies Nicolas Cage appears in; the y axis will be the number of people who drown after falling into in a swimming pool.  Put three things on the plot: a scatter plot showing the **real** relationship between the number of movies Nicolas Cage appears in and the number of people who drowned after falling into in a swimming pool on the y axis; and two lines of best fit for the predictions that each linear regression model predicted.


<!-- grade_id: cell-4d5668fff441e3a3, points: 1, grade: true, solution: true -->
```python
### BEGIN SOLUTION
dpred1 = reg.predict(n_films)
dpred2 = las.predict(n_films)

plt.scatter(n_films, pool_deaths)
plt.plot(n_films, dpred1)
plt.plot(n_films, dpred2)

plt.legend(["LLS", "Lasso"])
plt.xlabel("Nic Cage Films")
plt.ylabel("Deaths by Drowning")
plt.show()
### END SOLUTION
```

(iv) (2 points) Print out the R2 score for each model using the method shown in class to 2 decimal places. Which model performs better?  Does this answer the question whether Nicolas Cage is a national treasure or threat to life?

<!-- grade_id: cell-25b0d943244761db, points: 2, grade: true, solution: true -->
```python
### BEGIN SOLUTION
from sklearn.metrics import r2_score

print("Score LR: ", r2_score(pool_deaths, dpred1))
print("Score Lasso: ", r2_score(pool_deaths, dpred2))
print("I think the model that performs better is the LR model because it has a slightly higher R2 score, although they are very close to each other.")
print("I conclude that there is insufficient evidence that Nicolas Cage causes pool deaths because the R2 score is small.")
### END SOLUTION

# Fill in the blanks:
print("I think the model that performs better is... because ...")
print("I conclude that... because...")
```

#### Question 1b: Nicolas Cage: Good or Bad? [2 points]

Although it is clear from 1a that Nicolas Cage might possibly be indirectly responsible for people falling into pools and consequently dying, there is also a strong negative correlation between the number of films he stars in and the number of deaths from helicopter accidents. That is, the more movies Nicolas Cage starred in within a given year, the fewer people died in tragic helicopter accidents. So, we have to consider that he *might* save more lives than he ends.

The file `cage.csv` contains another column that represents the number of helicopter accidents that ended up killing the occupant. Use this data to answer the question that drove Abed to madness in Community: is Nicolas Cage good or bad?


<ol>
<li> Create two linear models (using `Linear Regression`) to predict the number of deaths both by helicopter accident and drowning (after falling into a pool) as a function of the number of films Nicolas Cage is in.
<li> Use these models to predict how many people will die by both helicopter accident and drowning (after falling into a pool) if Nicolas Cage stars in 6 movies in a year? (round to the nearest whole number).
<li> How many movies does Nicolas Cage have to star in if he wants to start saving lives? Specifically, how many movies does he need to star in for `(number of deaths by drowning) + (number of deaths by helicopter)` to be a negative number?
<ol>

Format your code in a function called `cage`. Your function should take no arguments, and return a list with three elements: `[a, b, c]`<br>
`a`: Number of people who would die by helicopter accident if Nicolas Cage stars in 6 movies, rounded to the nearest whole number.

`b`: Number of people who would die by drowning in a pool (after falling in, of course) if Nicolas Cage stars in 6 movies, rounded to the nearest whole number.

`c`: Minimum number of movies Nicolas Cage needs to star in for `number_of_people_dying_in_pools + number_of_people_dying_in_helicopters` to become a negative number.


<!-- grade_id: cell-f1fe5e5704b99405, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def cage():
    from sklearn import linear_model as lm

    n_films = []
    pool_deaths = []
    heli_deaths = []
    with open('cage.csv', 'r') as fh:
        reader = csv.DictReader(fh)

        for row in reader:
            n_films.append(int(row['number_of_nic_cage_movies']))
            pool_deaths.append(int(row['deaths_by_pool']))
            heli_deaths.append(int(row['deaths_by_helicopter']))
    n_films = [[i] for i in n_films]

    pool = lm.LinearRegression()
    pool.fit(n_films, pool_deaths)

    heli = lm.LinearRegression()
    heli.fit(n_films, heli_deaths)

    return (int(heli.predict(6)[0]), int(pool.predict(6)[0]), 58)
### END SOLUTION    
```

```python
# Test case for number of people dying by dying by helicopter
assert cage()[0] == 24
```

```python
# Test case to make sure your list is the correct length
assert len(cage()) == 3
```

<!-- grade_id: cell-6c22b20f3b706467, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert cage()[1] == 122
### END HIDDEN TESTS
```

<!-- grade_id: cell-a05d3b6ec9e90dee, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert cage()[2] == 58
### END HIDDEN TESTS
```

*In case you are wondering where we came up with these correlations, check out the website Spurious Correlations at http://www.tylervigen.com/spurious-correlations.  It contains such interesting facts as:*

- *the per capita consumption of cheese in the United States is strongly correlated with the number of people who died by becoming tangled in their bedsheets (don't eat cheese in bed!)*
- *the total revenue generated by video game arcades is strongly correlated with the number of computer science doctorates awarded in the US (nerds like video games!)*

### Question 2: Does Donald Trump's Mood Influence the Stock Market? [12 points]

The rest of the lab is designed to answer the above question by breaking it into a series of smaller steps, using the tools you've learnt so far.

#### Part a: [4 points]
Using the file `http://joey.legere.io/files/trump_lab10.json`, create a dictionary that maps each day (represeted as a datetime object) to a sentiment score for all tweets from that day.

Hints:
* Use the `time` field for each tweet to create a datetime structure. Crop the time so that it only contains the date. That is, if the text is `2018-03-05T12:57:12.000Z`, create a datetime structure using only the first few characters: `2018-03-05`.  (We are not worrying about timezones.)
* Calculate the sentiment analysis score for each date using VADER from nltk. If the output from vader is the variable s, the sentiment score of a tweet would just be s['pos'] - s['neg']. For simplicity, you do not need to average the scores based on the number of tweets in a given day, just sum the sentiments of all tweets that day.  Save this in a dictionary called `sents` with the keys being the datetime object.

```python
# This will ensure that the VADER lexicon datafile is installed before running your solution.

import nltk
nltk.download('vader_lexicon')
```

<!-- grade_id: cell-e8a157df892adb04, points: 2, grade: true, solution: true -->
```python
sents = {}

### BEGIN SOLUTION
import urllib.request
import json
import datetime
import nltk.sentiment.vader

url = 'http://joey.legere.io/files/trump_lab10.json'
response = urllib.request.urlopen(url)
charset = response.info().get_content_charset() or 'ascii'
s = response.read().decode(charset)
j = json.loads(s)

sents = {}
vader = nltk.sentiment.vader.SentimentIntensityAnalyzer()
for tweet in j:
    d = datetime.datetime.strptime(tweet["time"][:10], "%Y-%m-%d")
    s = vader.polarity_scores(tweet["text"])
    if d in sents:
        sents[d] += s['pos'] - s['neg']
    else:
        sents[d] = s['pos'] - s['neg']
### END SOLUTION
```

If you want to check your work so far, we computed that Trump's sentiment on 2018-02-14 was 0.16499999999999998.  (Happy Valentine's Day, I guess?)

```python
assert round(sents[datetime.datetime(2018, 2, 14)], 3) == 0.165
```

<!-- grade_id: cell-05e11d07cf88103c, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert round(sents[datetime.datetime(2018, 2, 21)], 2) == 1.33
### END HIDDEN TESTS
```

#### Part b: [1 point]

Using the file `DJI.csv` included in the assignment on JupyterHub, create a dictionary that maps each day (represeted as a datetime object) to the overall change in the stock market for that day, namely the closing price minus the opening price.  Save this in a dictionary called `pricechanges` with the keys being the datetime object.

<!-- grade_id: cell-913a040979cb2c29, points: 0, grade: false, solution: true -->
```python
pricechanges = {}

### BEGIN SOLUTION
import csv

pricechanges = {}
with open('DJI.csv', 'r') as fh:
    reader = csv.DictReader(fh)
    for r in reader:
        d = datetime.datetime.strptime(r["Date"], "%Y-%m-%d")
        s = float(r["Close"]) - float(r["Open"])
        pricechanges[d] = s
### END SOLUTION
```

If you want to check your work so far, we computed that the stock market change on 2018-02-14 was 357.67.

```python
assert round(pricechanges[datetime.datetime(2018, 2, 14)], 2) == 357.67
```

<!-- grade_id: cell-bf3b2eabbcdce3ff, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert round(pricechanges[datetime.datetime(2018, 2, 21)], 2) == -190.28
### END HIDDEN TESTS
```

#### Part c: [1 point]

We will only try to find a correlation on days when BOTH Trump tweeted AND the stock market was open.  Create three lists:

1. A list called `dates` containing the dates when BOTH Trump tweeted and the stock market was open.
2. A list containing Trump's tweets' sentiments on those dates, in the same order as the first list.
3. A list containing the stock market price difference on those dates, in the same order as the first list.

<!-- grade_id: cell-789909eda1d1f809, points: 1, grade: true, solution: true -->
```python
dates = []

### BEGIN SOLUTION
dates = []
sents_by_date = []
pricechanges_by_date = []
for date in sents:
    if date in pricechanges:
        dates.append(date)
        sents_by_date.append(sents[date])
        pricechanges_by_date.append(pricechanges[date])
### END SOLUTION
```

If you want to check your work, here are three facts:

- The stock market was not open on February 25, so that should not be in your list of dates.
- The stock market was open on February 26, but Trump didn't tweet that day, so that should not be in your list of dates.
- Trump tweeted and the stock market was open on February 27, so that should be in your list of dates.

```python
assert datetime.datetime(2018, 2, 25) not in dates
assert datetime.datetime(2018, 2, 26) not in dates
assert datetime.datetime(2018, 2, 27) in dates
```

#### Part d: [1 point]

Create three plots:

1. A line plot, showing dates on the x axis and Trump's sentiment on the y axis.
2. A line plot, showing dates on the x axis and stock price changes on the y axis.
3. A scatter plot, showing Trump's sentiment on the x axis and stock price changes on the y axis.

<!-- grade_id: cell-a61a91d5300cdcee, points: 1, grade: true, solution: true -->
```python
### BEGIN SOLUTION
import matplotlib.pyplot as plt

%matplotlib inline

plt.plot(dates, sents_by_date)
plt.xlabel("Date")
plt.ylabel("Trump's sentiment")
plt.show()

plt.plot(dates, pricechanges_by_date)
plt.xlabel("Date")
plt.ylabel("Price changes by date")
plt.show()

plt.scatter(sents_by_date, pricechanges_by_date)
plt.xlabel("Trump's sentiment")
plt.ylabel("Price changes")
plt.show()
### END SOLUTION
```

#### Part d: Linear regression [1 point]

Apply linear regression to develop a model for predicting the stock market price change based on Trump's sentiment.

1. Print the coeffient and intercept calculated by the regression model.
2. Then use the linear regression model to predict what the stock market price change would be based on Trump's sentiment.  (We will commit a statistical no-no: we didn't split the data into a training set and a testing set.  Instead, we'll just do predictions on the training data.)

<!-- grade_id: cell-51c15eb895b62653, points: 1, grade: true, solution: true -->
```python
### BEGIN SOLUTION
from sklearn import linear_model

regr = linear_model.LinearRegression()
sents_by_date_fixed = fix1Darray(sents_by_date)
regr.fit(sents_by_date_fixed, pricechanges_by_date)

print('Estimated coefficient: ', regr.coef_)
print('Estimated intercept:   ', regr.intercept_)

pred_pricechanges = regr.predict(sents_by_date_fixed)
### END SOLUTION
```

### Part e: Evaluation [2 points]

Create a plot just like the scatter plot in part (c) above, but add in a line of best fit using the predicted data.

Print the R2 score.  

Based on the R2 score, how would you answer the question: Are stock market prices affected by President Trump's mood?

<!-- grade_id: cell-da37d118c647ba55, points: 2, grade: true, solution: true -->
```python
### BEGIN SOLUTION
plt.scatter(sents_by_date, pricechanges_by_date)
plt.plot(sents_by_date, pred_pricechanges)
plt.xlabel("Trump's sentiment")
plt.ylabel("Price changes")
plt.show()

from sklearn.metrics import r2_score

r2score = r2_score(pricechanges_by_date, pred_pricechanges)
print("R2 score: {:.2f}".format(r2score))

print("I conclude that this experiment does not given any evidence that the stock market is affected by President Trump's mood, because the R2 score is small.  It is possible that prices are affected by his mood in some other way, only that this experiment has not shown they are related in this exact way.")
### END SOLUTION

# Fill in the blanks
print("I conclude that... because...")
```

#### Part f: Bonus (2 marks)

Now let's abuse statistics, in the vein of Spurious Correlations.  For a bonus of up to 2 marks, do some additional experiment on whether Trump's mood affects the stock market.  Maybe use a different type of regression?  Maybe use some other way of measuring his sentiment?  Maybe check another stock market index or stock?  Perhaps Trump's use of the word "Obama" correlates with price changes for the OncBioMune Pharmaceuticals, Inc. corporation (whose stock symbol is OBMP)?  Full marks if you demonstrate a correlation with an R2 score of at least 0.95.

<!-- grade_id: cell-289f3011584a984a, points: 2, grade: true, solution: true -->
```python
# YOUR CODE HERE
```