## COMP SCI 1MD3 Introduction to Programming, Winter 2018
### Douglas Stebila (Instructor), Joey Legere, Karl Knopf, Natalie Chin, Victor Chen (TAs)
### Lab 8 Assigned Saturday, March 10, Due Friday March 16, 5pm
### Maximum grade: 20 / 16

The purpose of this lab is to:
* Work with external files
* Learn about csv and JSON data
* Work with network / remote files

#### Practice Question 1: Reading and Writing Plaintext Files

One of the most common ways to represent data is as *plain text*. This is different from rich text. as used by programs like Word, or your web browser, which can be stylized in some way, such as **bold** *italicized*, or <font color='red'>coloured</font> . Plain text data is just that--plain text, with no styling.

Write a function `copyfile(infile, outfile)` that, given two files names: `infile` and `outfile`, reads the contents of `infile`, prints them on the screen, and then writes them into `outfile`.

```python
# Your code here

```

The cell below should print out the contents of the file if you've read it correctly:

```python
copyfile("practice1.txt", "practice1_out.txt")
```

Output:
```
Woo! It worked!
```

The cell below is a Jupyter magic command that tries to open `practice1_out.txt` and display it. If you wrote your file correctly, this should display the same as above:

```python
% cat practice1_out.txt
```

Output:
```
Woo! It worked!
```

#### Practice Question 2: Working with CSV files

A comma-separated value (csv) file is a very common way to represent a spreadsheet in a plain text file. In a csv file, each value is separated by a comma, and each row is separated by a newline character. For instance, the following table:

|name|age|salary|
|-|-||-|
|Odie|23|19000|
|Garfield|25|22000|
|John|22|21000|

Is not meant to be a commentary on the economy at all. It is simply an example of a table that could be represented in a csv file like so:

name,age,salary
Odie,23,19000
Garfield,25,22000
John,22,21000

Write a program `getrow(infile, n)` that reads a .csv file, `infile`, and prints out the column labels, and the `nth` row of `infile`. For instance, in the above example, calling `getcolumn(infile, 1)` would print out:
```
['name', 'age', 'salary']
['Odie', '23', '19000']
```

You may wish to use Python's built in <a href=https://docs.python.org/3/library/csv.html>csv library</a> for this task.

```python
# Your code here.

```

```python
getrow('practice2.csv', 1)
```

Output:
```
['name', 'age', 'salary']
['Odie', '23', '19000']
```

This should produce the output
```
['name', 'age', 'salary']
['Odie', '23', '19000']
```

#### Practice Question 3: Network Files and JSON

JavaScript Object Notation (JSON) data is another common way of representing more structed data. In most cases, you can treat JSON data as a list of Python dictionaries. 

Often a JSON object will have many fields for each entry. Consider the example below, which organizes a few books on JavaScript:

```python
BOOKS = \
[
    {
        "title": "Professional JavaScript",
        "authors": [
            "Nicholas C. Zakas"
        ],
        "edition": 3,
        "year": 2011
    },
    {
        "title": "Professional JavaScript",
        "authors": [
            "Nicholas C. Zakas"
        ],
        "edition": 2,
        "year": 2009
    },
    {
        "title": "Professional Ajax",
        "authors": [
            "Nicholas C. Zakas",
            "Jeremy McPeak",
            "Joe Fawcett"
        ],
        "edition": 2,
        "year": 2008
    }
]
```

Getting the first author of a book in the list isn't too challenging:

```python
for i in range(len(BOOKS)):
    
    print(BOOKS[i]["authors"][0])
```

Output:
```
Nicholas C. Zakas
Nicholas C. Zakas
Nicholas C. Zakas
```

Write a function `remotejson(url, fields)` that reads in a JSON object from a remote url, and prints out every object from the given fields. Assume we are only interested in the top level fields--that is, you only need to index `OBJ[i][field]`, and no subfields that it might contain. However, if more than one field is given, return their data side by side.

You may wish to separate your output with an extra newline between outputs in order to make it easier to read.

Note: Use the built in `urllib` library to read the files from the remote server, and `json` in order to parse it.

```python
# Your code here.

## ANSWER, remove

```

The cell below, when run correctly, should print a list of some of Dr. Stebila's favourite songs, and where to listen to them.

```python
remotejson("http://brain.mcmaster.ca/joey/dr_stebilas_songs.json", ['Title', "Artist", "URL"])
```

Output:
```
Baby
Justin Bieber
https://www.youtube.com/watch?v=kffacxfA7G4

What Makes You Beautiful
One Direction
https://www.youtube.com/watch?v=QJO3ROT-A4E

Friday
Rebecca Black
https://www.youtube.com/watch?v=kfVsfOSbJY0

Barbie Girl
Aqua
https://www.youtube.com/watch?v=ZyhrYis509A
```

*[Comment from Dr. Stebila: Joey only guessed 2 of my favourite songs.]*

#### Question 1a: I know words [3 points]

During his presidential campaing in 2016, Donald Trump said “I know words. I have the best words.”

Let's explore this idea in an objective, scientific manner.

Write a function `tweets_per_day(jsonurl)` that, given the website address (URL) of a JSON file, counts the number of tweets that are sent on every day of the week. Return your result as a dictionary, mapping the day of the week to the total number of tweets sent on that day.

For this question, you may want to open the file `http://brain.mcmaster.ca/joey/trump_tweets.json` and manually inspect the fields to help parse the data.
    

<!-- grade_id: cell-a952ea8b4c173603, points: 0, grade: false, solution: true -->
```python
import json
import urllib.request
# Write your answer here.

### BEGIN SOLUTION
import json
import urllib.request

def tweets_per_day(jsonurl):
    
    req = urllib.request.urlopen(jsonurl)
    charset = req.info().get_content_charset() or 'ascii'
    s = req.read().decode(charset)
    tweets = json.loads(s)

    tpd = {}
    
    for t in tweets:
        
        try:
            tpd[t['day']] += 1
        except KeyError:
            tpd[t['day']] = 1
    
    return tpd
### END SOLUTION
```

```python
tweets_per_day('http://brain.mcmaster.ca/joey/trump_tweets.json') == {
 'Friday': 33,
 'Monday': 19,
 'Saturday': 35,
 'Sunday': 26,
 'Thursday': 30,
 'Tuesday': 28,
 'Wednesday': 29
}
```

Output:
```
True
```

<!-- grade_id: cell-9e852f206f21e307, points: 1, grade: true, solution: false -->
```python
# Hidden test for Q1a
### BEGIN HIDDEN TESTS
assert tweets_per_day('http://legere.io/files/trump_1.json') == {'Friday': 26,
 'Monday': 14,
 'Saturday': 31,
 'Sunday': 25,
 'Thursday': 30,
 'Tuesday': 28,
 'Wednesday': 29}
### END HIDDEN TESTS
```

<!-- grade_id: cell-af2bb0f99332df06, points: 1, grade: true, solution: false -->
```python
# Hidden test for Q1a
### BEGIN HIDDEN TESTS
assert tweets_per_day('http://legere.io/files/trump_2.json') == {'Friday': 24,
 'Monday': 14,
 'Saturday': 31,
 'Sunday': 25,
 'Thursday': 24,
 'Tuesday': 27,
 'Wednesday': 19}
### END HIDDEN TESTS
```

<!-- grade_id: cell-f9ae7ae1dfb9b294, points: 1, grade: true, solution: false -->
```python
# Hidden test for Q1a
### BEGIN HIDDEN TESTS
assert tweets_per_day('http://legere.io/files/trump_3.json') == {'Friday': 24,
 'Monday': 14,
 'Saturday': 27,
 'Sunday': 23,
 'Thursday': 24,
 'Tuesday': 20,
 'Wednesday': 19}
### END HIDDEN TESTS
```

#### Question 1b: The Best Words [4 points]

Write a function `count_occurences(jsonurl, word)` that counts the number of times that `word` occurs across all of the given tweets. This should be case insensitive, but for simplicity, you only need to match whole words. 

Make sure you match whole words only, and not substrings. For example, if `word = "dont"`, then "dont", "DONT", and "dOnT" would all count, but "perioDONTal" should not. Remove all punctuation in the text of the tweet (if there is any). "don't" should be correctly counted for `word="dont"`.

```python
from string import punctuation
print("Remove all of these symbols from the tweets: ", punctuation)
```

Output:
```
Remove all of these symbols from the tweets:  !"#$%&'()*+,-./:;<=>?@[\]^_`{|}~
```

<!-- grade_id: cell-e9b4c034335292a9, points: 0, grade: false, solution: true -->
```python
import json
import urllib.request
from string import punctuation
# Write your answer here

### BEGIN SOLUTION
import json
import urllib.request
from string import punctuation

def count_occurences(jsonurl, word):
    
    req = urllib.request.urlopen(jsonurl)
    charset = req.info().get_content_charset() or 'ascii'
    s = req.read().decode(charset)
    tweets = json.loads(s)
        
    count = 0
    word = word.upper()
    for t in tweets:
        
        clean = ''
        for char in t['text']:
            if char not in punctuation: clean += char
        clean = clean.upper()
        clean = clean.split(" ")

        count += clean.count(word)

    return count
### END SOLUTION
```

```python
count_occurences('http://brain.mcmaster.ca/joey/trump_tweets.json', 'crooked') == 5
```

Output:
```
True
```

```python
count_occurences('http://brain.mcmaster.ca/joey/trump_tweets.json', 'fake') == 9
```

Output:
```
True
```

```python
count_occurences('http://brain.mcmaster.ca/joey/trump_tweets.json', 'obama') == 18
```

Output:
```
True
```

```python
count_occurences('http://brain.mcmaster.ca/joey/trump_tweets.json', 'wall') == 8
```

Output:
```
True
```

<!-- grade_id: cell-6aa804bc80448063, points: 1, grade: true, solution: false -->
```python
# Hidden test for Q1b
### BEGIN HIDDEN TESTS
assert count_occurences('http://legere.io/files/trump_1.json', 'crooked') == 4
### END HIDDEN TESTS
```

<!-- grade_id: cell-a35dfcaf05cdda6f, points: 1, grade: true, solution: false -->
```python
# Hidden test for Q1b
### BEGIN HIDDEN TESTS
assert count_occurences('http://legere.io/files/trump_2.json', 'wall') == 5
### END HIDDEN TESTS
```

<!-- grade_id: cell-36d596ab96c5d78d, points: 1, grade: true, solution: false -->
```python
# Hidden test for Q1b
### BEGIN HIDDEN TESTS
assert count_occurences('http://legere.io/files/trump_2.json', 'obama') == 15
### END HIDDEN TESTS
```

<!-- grade_id: cell-62467bb65233149b, points: 1, grade: true, solution: false -->
```python
# Hidden test for Q1b
### BEGIN HIDDEN TESTS
assert count_occurences('http://legere.io/files/trump_3.json', 'russia') == 12
### END HIDDEN TESTS
```

#### Question 1c: The Most Words. The Best Words. [3 points]

Write a function `words_per_day(jsonurl, word)` that counts the total number of times that `word` is tweeted versus the day of the week.

<!-- grade_id: cell-495893b89823d6bc, points: 0, grade: false, solution: true -->
```python
import json
import urllib.request
from string import punctuation

# Write your answer here.


### BEGIN SOLUTION
import json
import urllib.request
from string import punctuation

def words_per_day(jsonurl, word):
    
    req = urllib.request.urlopen(jsonurl)
    charset = req.info().get_content_charset() or 'ascii'
    s = req.read().decode(charset)
    tweets = json.loads(s)
        
    wpd = {}
    word = word.upper()
    
    
    for t in tweets:
        
        clean = ''
        for char in t['text']:
            if char not in punctuation: clean += char
        clean = clean.upper()
        clean = clean.split(" ")

        try:
            wpd[t['day']] += clean.count(word)
        except KeyError:
            wpd[t['day']] = clean.count(word)


    return wpd
### END SOLUTION
```

```python
words_per_day('http://brain.mcmaster.ca/joey/trump_tweets.json', 'america') == {
 'Friday': 1,
 'Monday': 3,
 'Saturday': 0,
 'Sunday': 2,
 'Thursday': 2,
 'Tuesday': 1,
 'Wednesday': 2
}
```

Output:
```
True
```

```python
words_per_day('http://brain.mcmaster.ca/joey/trump_tweets.json', 'china') == {
 'Friday': 0,
 'Monday': 0,
 'Saturday': 0,
 'Sunday': 1,
 'Thursday': 1,
 'Tuesday': 0,
 'Wednesday': 0
}
```

Output:
```
True
```

```python
words_per_day('http://brain.mcmaster.ca/joey/trump_tweets.json', 'great') == {
 'Friday': 8,
 'Monday': 4,
 'Saturday': 1,
 'Sunday': 8,
 'Thursday': 8,
 'Tuesday': 8,
 'Wednesday': 9
}
```

Output:
```
True
```

```python
words_per_day('http://brain.mcmaster.ca/joey/trump_tweets.json', 'hillary') == {
 'Friday': 2,
 'Monday': 0,
 'Saturday': 0,
 'Sunday': 1,
 'Thursday': 0,
 'Tuesday': 1,
 'Wednesday': 0
}
```

Output:
```
True
```

<!-- grade_id: cell-655a40239a386869, points: 1, grade: true, solution: false -->
```python
# Hidden test for Q1c
### BEGIN HIDDEN TESTS
assert words_per_day('http://legere.io/files/trump_1.json', 'crooked') == {'Friday': 1,
 'Monday': 0,
 'Saturday': 0,
 'Sunday': 2,
 'Thursday': 0,
 'Tuesday': 1,
 'Wednesday': 0}
### END HIDDEN TESTS
```

<!-- grade_id: cell-fd9e9257ac15624e, points: 1, grade: true, solution: false -->
```python
# Hidden test for Q1c
### BEGIN HIDDEN TESTS
assert words_per_day('http://legere.io/files/trump_2.json', 'news') == {'Friday': 1,
 'Monday': 0,
 'Saturday': 3,
 'Sunday': 2,
 'Thursday': 1,
 'Tuesday': 2,
 'Wednesday': 4}
### END HIDDEN TESTS
```

<!-- grade_id: cell-a2fbd06ee45ab5ca, points: 1, grade: true, solution: false -->
```python
# Hidden test for Q1c
### BEGIN HIDDEN TESTS
assert words_per_day('http://legere.io/files/trump_3.json', 'great') == {'Friday': 7,
 'Monday': 3,
 'Saturday': 1,
 'Sunday': 7,
 'Thursday': 7,
 'Tuesday': 6,
 'Wednesday': 6}
### END HIDDEN TESTS
```

#### Question 2: Stocks (Again) [6 points]

Recall the `stocks` question from lab 4:

> Let's assume I'm playing the stock market - buy low, sell high. I'm a day trader, so I want to get in and out of a stock before the day is done, and I want to time my trades so that I make the biggest gain possible. And obviously I can't buy in the future and sell in the past.
>
> Write a function stocks(s) that emits the two trades in chronological order - what you think I should buy at and sell at - in order to maximize my profit.

Repeat this question, but load your data from a .csv file on a remote server. The .csv file will be formatted like so:
```
Date,Open,High,Low,Close,Adj Close,Volume
2013-12-24,16295.700195,16360.599609,16295.700195,16357.549805,16357.549805,33640000
2013-12-26,16370.969727,16483.000000,16370.969727,16479.880859,16479.880859,50160000
2013-12-27,16486.369141,16529.009766,16461.230469,16478.410156,16478.410156,47230000
2013-12-30,16484.509766,16504.349609,16476.869141,16504.289063,16504.289063,54220000
```

Assume that you buy at the value in the "Low" column, and sell at the value in "High". You may not buy and sell on the same day. Your functions, `stocks(s)` should take in a string of the filename of a CSV file containing the data. You may assume the file is in chronological order.

<!-- grade_id: cell-272ccd76b9909416, points: 0, grade: false, solution: true -->
```python
# Write your answer here.

### BEGIN SOLUTION
import csv
def stocks(csvfilename):
    
    with open(csvfilename, 'r') as csvfh:
    
        c = csv.DictReader(csvfh)
        data = []
        for line in c:
            data.append(line)

        maxprof = 0
        trades = [0, 0]

        for i in range(len(data)):
            data[i]["Low"] = float(data[i]["Low"])
            for j in range(i+1, len(data)):
                data[j]["High"] = float(data[j]["High"])
                if data[j]["High"] - data[i]["Low"] > maxprof:
                    maxprof = data[j]["High"] - data[i]["Low"]
                    trades = (data[i]["Low"], data[j]["High"])
                    
        return trades
### END SOLUTION    
```

This cell should print the values of the trades. They should be [15340.69043, 26616.710938]

```python
trades = stocks("DJI.csv")
print(trades)
```

Output:
```
(15340.69043, 26616.710938)
```

The cell below is how the autograder will be evaluating your answer. It only needs to be correct to two decimal places. Make sure you return a list of floating point numbers.

```python
assert round(trades[0], 2) == 15340.69
assert round(trades[1], 2) == 26616.71
```

<!-- grade_id: cell-ad372f0ca2ab122d, points: 3, grade: true, solution: false -->
```python
# Hidden test for Q2
### BEGIN HIDDEN TESTS
trades = stocks("DJI2.csv")
assert round(trades[0], 2) == 15370.33
assert round(trades[1], 2) == 26246.19
### END HIDDEN TESTS
```

<!-- grade_id: cell-b58cb12369ce57ba, points: 3, grade: true, solution: false -->
```python
# Hidden test for Q2
### BEGIN HIDDEN TESTS
trades = stocks("DJI3.csv")
assert round(trades[0], 2) == 15450.56
assert round(trades[1], 2) == 25439.78
### END HIDDEN TESTS
```

#### Bonus: Working with Pillow (4 points)

The Python Imaging Library (PIL or pillow) is a library for loading and manipulating images in Python. The below example creates an image that is black in the top left, becomes more red as you move to the right, and more green as you move towards the bottom.

```python
from PIL import Image

# PIL accesses images in Cartesian co-ordinates, so it is Image[columns, rows]
img = Image.new( 'RGB', (250,250), "black") # create a new black image
pixels = img.load() # create the pixel map

for i in range(img.size[0]):    # for every col:
    for j in range(img.size[1]):    # For every row
        pixels[i,j] = (i, j, 0) # set the colour accordingly

img.show()
```

`img.show()` will try to show you the image, but it does not work on JupyterHub.  You will have to run this code on your local installation of Jupyter to work on your solution to this problem.  

But even on your local installation of Jupyter you may still also have to install PIL.  You can install PIL on your computer by running one of the following commands in your command-line terminal:

- `pip3 install pillow`
- `python3 -m pip install pillow`
- `sudo pip3 install pillow`

Each pixel is represented using an RGB value--that is, how much red, green, and blue there is in each pixel. The line:

`pixels[i,j] = (i, j, 0) # set the colour accordingly`

Accomplishes this by setting the pixel at coordinates (i, j) to have a red value of i (its x coordinate), a green value of j (its y coordinate), and a blue value of 0. Indexing works such that (0, 0) is at the top left corner.

This question will be manually graded. You may choose one of the following ideas, or implement your own idea using pillow (if it is sufficiently technically challenging enough, and most importantly, not copied from stackoverflow...)

* Write a function that takes in the path to an image, and darkens or lightens it by a set amount.
* Write a function that generates an image of your choosing (not the one from the example..)
* Write a function that generates your absolute best sketch of Donald Trump, programmatically.

If you use any outside sources for this question, add links to them in your comments.

<!-- grade_id: cell-8bae9d13fa3ace28, points: 4, grade: true, solution: true -->
```python
# BONUS
from PIL import Image
# Write your answer here
```